A Deep Learning Pipeline for Mapping in situ Network-level Neurovascular Coupling in Multi-photon Fluorescence Microscopy

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public review):

Summary:

In the manuscript the authors describe a new pipeline to measure changes in vasculature diameter upon optogenetic stimulation of neurons. The work is useful to better understand the hemodynamic response on a network /graph level.

Strengths:

The manuscript provides a pipeline that allows to detect changes in the vessel diameter as well as simultaneously allows to locate the neurons driven by stimulation.

The resulting data could provide interesting insights into the graph level mechanisms of regulating activity dependent blood flow.

Weaknesses:

(1) The manuscript contains (new) wrong statements and (still) wrong mathematical formulas.

The symbols in these formulas have been updated to disambiguate them, and the accompanying statements have been adjusted for clarity.

(2) The manuscript does not compare results to existing pipelines for vasculature segmentation (opensource or commercial). Comparing performance of the pipeline to a random forest classifier (illastik) on images that are not preprocessed (i.e. corrected for background etc.) seems not a particularly useful comparison.

We’ve now included comparisons to Imaris (a commercial) for segmentation and VesselVio (open-source) for graph extraction software.

For the ilastik comparison, the images were preprocessed prior to ilastik segmentation, specifically by doing intensity normalization.

Example segmentations utilizing Imaris have now been included. Imaris leaves gaps and discontinuities in the segmentation masks, as shown in Supplementary Figure 10. The Imaris segmentation masks also tend to be more circular in cross-section despite irregularities on the surface of the vessels observable in the raw data and identified in manual segmentation. This approach also requires days to months to generate per image stack.

“Comparison with commercial and open-source vascular analysis pipelines

To compare our results with those achievable on these data with other pipelines for segmentation and graph network extraction, we compared segmentation results qualitatively with Imaris version 9.2.1 (Bitplane) and vascular graph extraction with VesselVio [1]. For the Imaris comparison, three small volumes were annotated by hand to label vessels. Example slices of the segmentation results are shown in Supplementary Figure 10. Imaris tended to either over- or under-segment vessels, disregard fine details of the vascular boundaries, and produce jagged edges in the vascular segmentation masks. In addition to these issues with segmentation mask quality, manual segmentation of a single volume took days for a rater to annotate. To compare to VesselVio, binary segmentation masks (one before and one after photostimulation) generated with our deep learning models were loaded into VesselVio for graph extraction, as VesselVio does not have its own method for generating segmentation masks. This also facilitates a direct comparison of the benefits of our graph extraction pipeline to VesselVio. Visualizations of the two graphs are shown in Supplementary Figure 11. Vesselvio produced many hairs at both time points, and the total number of segments varied considerably between the two sequential stacks: while the baseline scan resulted in 546 vessel segments, the second scan had 642 vessel segments. These discrepancies are difficult to resolve in post-processing and preclude a direct comparison of individual vessel segments across time. As the segmentation masks we used in graph extraction derive from the union of multiple time points, we could better trace the vasculature and identify more connections in our extracted graph. Furthermore, VesselVio relies on the distance transform of the user supplied segmentation mask to estimate vascular radii; consequently, these estimates are highly susceptible to variations in the input segmentation masks.We repeatedly saw slight variations between boundary placements of all of the models we utilized (ilastik, UNet, and UNETR) and those produced by raters. Our pipeline mitigates this segmentation method bias by using intensity gradient-based boundary detection from centerlines in the image (as opposed to using the distance transform of the segmentation mask, as in VesselVio).”

(3) The manuscript does not clearly visualize performance of the segmentation pipeline (e.g. via 2d sections, highlighting also errors etc.). Thus, it is unclear how good the pipeline is, under what conditions it fails or what kind of errors to expect.

On reviewer’s comment, 2D slices have been added in the Supplementary Figure 4.

(4) The pipeline is not fully open-source due to use of matlab. Also, the pipeline code was not made available during review contrary to the authors claims (the provided link did not lead to a repository). Thus, the utility of the pipeline was difficult to judge.

All code has been uploaded to Github and is available at the following location: https://github.com/AICONSlab/novas3d

The Matlab code for skeletonization is better at preserving centerline integrity during the pruning of hairs from centerlines than the currently available open-source methods.

- Generalizability: The authors addressed the point of generalizability by applying the pipeline to other data sets. This demonstrates that their pipeline can be applied to other data sets and makes it more useful. However, from the visualizations it's unclear to see the performance of the pipeline, where the pipelines fails etc. The 3d visualizations are not particularly helpful in this respect . In addition, the dice measure seems quite low, indicating roughly 20-40% of voxels do not overlap between inferred and ground truth. I did not notice this high discrepancy earlier. A thorough discussion of the errors appearing in the segmentation pipeline would be necessary in my view to better assess the quality of the pipeline.

2D slices from the additional datasets have been added in the Supplementary Figure 13 to aid in visualizing the models’ ability to generalize to other datasets.

The dice range we report on (0.7-0.8) is good when compared to those (0.56-86) of 3D segmentations of large datasets in microscopy [2], [3], [4], [5], [6]. Furthermore, we had two additional raters segment three images from the original training set. We found that the raters had a mean inter class correlation of 0.73 [7]. Our model outperformed this Dice score on unseen data: Dice scores from our generalizability tests on C57 mice and Fischer rats on par or higher than this baseline.

Reviewer #2 (Public review):

The authors have addressed most of my concerns sufficiently. There are still a few serious concerns I have. Primarily, the temporal resolution of the technique still makes me dubious about nearly all of the biological results. It is good that the authors have added some vessel diameter time courses generated by their model. But I still maintain that data sampling every 42 seconds - or even 21 seconds - is problematic. First, the evidence for long vascular responses is lacking. The authors cite several papers:

Alarcon-Martinez et al. 2020 show and explicitly state that their responses (stimulus-evoked) returned to baseline within 30 seconds. The responses to ischemia are long lasting but this is irrelevant to the current study using activated local neurons to drive vessel signals.

Mester et al. 2019 show responses that all seem to return to baseline by around 50 seconds post-stimulus.

In Mester et al. 2019, diffuse stimulations with blue light showed a return to baseline around 50 seconds post-stimulus (cf. Figure 1E,2C,2D). However, focal stimulations where the stimulation light is raster scanned over a small region focused in the field of view show longer-lasting responses (cf. Figure 4) that have not returned to baseline by 70 seconds post-stimulus [8]. Alarcon-Martinez et al. do report that their responses return baseline within 30 seconds; however, their physiological stimulation may lead to different neuronal and vessel response kinetics than those elicited by the optogenetic stimulations as in current work.

O'Herron et al. 2022 and Hartmann et al. 2021 use opsins expressed in vessel walls (not neurons as in the current study) and directly constrict vessels with light. So this is unrelated to neuronal activity-induced vascular signals in the current study.

We agree that optogenetic activation of vessel-associated cells is distinct from optogenetic activation of neurons, but we do expect the effects of such perturbations on the vasculature to have some commonalities.

There are other papers including Vazquez et al 2014 (PMID: 23761666) and Uhlirova et al 2016 (PMID: 27244241) and many others showing optogenetically-evoked neural activity drives vascular responses that return back to baseline within 30 seconds. The stimulation time and the cell types labeled may be different across these studies which can make a difference. But vascular responses lasting 300 seconds or more after a stimulus of a few seconds are just not common in the literature and so are very suspect - likely at least in part due to the limitations of the algorithm.

The photostimulation in Vazquez et al. 2014 used diffuse photostimulation with a fiberoptic probe similar to Mester et al. 2019 as opposed to raster scanning focal stimulation we used in this study and in the study by Mester et al. 2019 where we observed the focal photostimulation to elicited longer than a minute vascular responses. Uhlirova et al. 2016 used photostimulation powers between 0.7 and 2.8 mW, likely lower than our 4.3 mW/mm2 photostimulation. Further, even with focal photostimulation, we do see light intensity dependence of the duration of the vascular responses. Indeed, in Supplementary Figure 2, 1.1 mW/mm2 photostimulation leads to briefer dilations/constrictions than does 4.3 mW/mm2; the 1.1 mW/mm2 responses are in line, duration wise, with those in Uhlirova et al. 2016.

Critically, as per Supplementary Figure 2, the analysis of the experimental recordings acquired at 3-second temporal resolution did likewise show responses in many vessels lasting for tens of seconds and even hundreds of seconds in some vessels.

Another major issue is that the time courses provided show that the same vessel constricts at certain points and dilates later. So where in the time course the data is sampled will have a major effect on the direction and amplitude of the vascular response. In fact, I could not find how the "response" window is calculated. Is it from the first volume collected after the stimulation - or an average of some number of volumes? But clearly down-sampling the provided data to 42 or even 21 second sampling will lead to problems. If the major benefit to the field is the full volume over large regions that the model can capture and describe, there needs to be a better way to capture the vessel diameter in a meaningful way.

In the main experiment (i.e. excluding the additional experiments presented in the Supplementary Figure 2 that were collected over a limited FOV at 3s per stack), we have collected one stack every 42 seconds. The first slice of the volume starts following the photostimulation, and the last slice finishes at 42 seconds. Each slice takes ~0.44 seconds to acquire. The data analysis pipeline (as demonstrated by the Supplementary Figure 2) is not in any way limited to data acquired at this temporal resolution and - provided reasonable signal-to-noise ratio (cf. Figure 5) - is applicable, as is, to data acquired at much higher sampling rates.

It still seems possible that if responses are bi-phasic, then depth dependencies of constrictors vs dilators may just be due to where in the response the data are being captured - maybe the constriction phase is captured in deeper planes of the volume and the dilation phase more superficially. This may also explain why nearly a third of vessels are not consistent across trials - if the direction the volume was acquired is different across trials, different phases of the response might be captured.

Alternatively, like neuronal responses to physiological stimuli, the vascular responses elicited by increases in neuronal activity may themselves be variable in both space and time.

I still have concerns about other aspects of the responses but these are less strong. Particularly, these bi-phasic responses are not something typically seen and I still maintain that constrictions are not common. The authors are right that some papers do show constriction. Leaving out the direct optogenetic constriction of vessels (O'Herron 2022 & Hartmann 2021), the Alarcon-Martinez et al. 2020 paper and others such as Gonzales et al 2020 (PMID: 33051294) show different capillary branches dilating and constricting. However, these are typically found either with spontaneous fluctuations or due to highly localized application of vasoactive compounds. I am not familiar with data showing activation of a large region of tissue - as in the current study - coupled with vessel constrictions in the same region. But as the authors point out, typically only a few vessels at a time are monitored so it is possible - even if this reviewer thinks it unlikely - that this effect is real and just hasn't been seen.

Uhlirova et al. 2016 (PMID: 27244241) observed biphasic responses in the same vessel with optogenetic stimulation in anesthetized and unanesthetized animals (cf Fig 1b and Fig 2, and section “OG stimulation of INs reproduces the biphasic arteriolar response”). Devor et al. (2007) and Lindvere et al. (2013) also reported on constrictions and dilations being elicited by sensory stimuli.

I also have concerns about the spatial resolution of the data. It looks like the data in Figure 7 and Supplementary Figure 7 have a resolution of about 1 micron/pixel. It isn't stated so I may be wrong. But detecting changes of less than 1 micron, especially given the noise of an in vivo prep (brain movement and so on), might just be noise in the model. This could also explain constrictions as just spurious outputs in the model's diameter estimation. The high variability in adjacent vessel segments seen in Figure 6C could also be explained the same way, since these also seem biologically and even physically unlikely.

Thank you for your comment. To address this important issue, we performed an additional validation experiment where we placed a special order of fluorescent beads with a known diameter of 7.32 ± 0.27um, imaged them following our imaging protocol, and subsequently used our pipeline to estimate their diameter. Our analysis converged on the manufacturer-specified diameters, estimating the diameter to be 7.34 ± 0.32. The manuscript has been updated to detail this experiment, as below:

Methods section insert

“Second, our boundary detection algorithm was used to estimate the diameters of fluorescent beads of a known radius imaged under similar acquisition parameters. Polystyrene microspheres labelled with Flash Red (Bangs Laboratories, inc, CAT# FSFR007) with a nominal diameter of 7.32um and a specified range of 7.32 ± 0.27um as determined by the manufacturer using a Coulter counter were imaged on the same multiphoton fluorescence microscope set-up used in the experiment (identical light path, resonant scanner, objective, detector, excitation wavelength and nominal lateral and axial resolutions, with 5x averaging). The images of the beads had a higher SNR than our images of the vasculature, so Gaussian noise was added to the images to degrade the SNR to the same level of that of the blood vessels. The images of the beads were segmented with a threshold, centroids calculated for individual spheres, and planes with a random normal vector extracted from each bead and used to estimate the diameter of the beads. The same smoothing and PSF deconvolution steps were applied in this task. We then reported the mean and standard deviation of the distribution of the diameter estimates. A variety of planes were used to estimate the diameters.”

Results Section Insert

“Our boundary detection algorithm successfully estimated the radius of precisely specified fluorescent beads. The bead images had a signal-to-noise ratio of 6.79 ± 0.16 (about 35% higher than our in vivo images): to match their SNR to that of in vivo vessel data, following deconvolution, we added Gaussian noise with a standard deviation of 85 SU to the images, bringing the SNR down to 5.05 ± 0.15. The data processing pipeline was kept unaltered except for the bead segmentation, performed via image thresholding instead of our deep learning model (trained on vessel data). The bead boundary was computed following the same algorithm used on vessel data: i.e., by the average of the minimum intensity gradients computed along 36 radial spokes emanating from the centreline vertex in the orthogonal plane. To demonstrate an averaging-induced decrease in the uncertainty of the bead radius estimates on a scale that is finer than the nominal resolution of the imaging configuration, we tested four averaging levels in 289 beads. Three of these averaging levels were lower than that used on the vessels, and one matched that used on the vessels (36 spokes per orthogonal plane and a minimum of 10 orthogonal planes per vessel). As the amount of averaging increased, the uncertainty on the diameter of the beads decreased, and our estimate of the bead's diameter converged upon the manufacturer's Coulter counter-based specifications (7.32 ± 0.27um), as tabulated in Table 1.”

Reviewer #1 (Recommendations for the authors):

Comments to the authors replies to the reviews:

- Supplementary Figure 13:

As indicated before the 3d images + scale makes it impossible to judge the quality of the outputs.

As aforementioned, 2D slices have been added to the Supplementary Figure 13.

- Supplementary Table 3:

There is a significant increase in the Hausdorrf and Mean Surface Distance measures for the new data, why ?

A single vessel being missed by either the rater or the model would significantly affect the Hausdorff distance (HD) and by extension Mean Surface Distance: this is particularly pertinent in the LSFM image with its much larger FOV and thus a potential for much larger max distances to result from missed vessels in the prediction or ground truth data. Large Hausdorff distances may indicate a vessel was missed in either the ground truth or the segmentation mask.

Of note, a different rater annotated these additional datasets from the raters labeling the ground truth data. There is a high variability in boundary placements between raters. On a test where three raters segmented the same three images from the original dataset, we computed a ICC of 0.73 across their segmentations. Our model Dice scores on predictions in out-of-distribution data sets were on par with the inter-rater ICC on the Thy1ChR2 2PFM data.

- Supplementary Figure 2: The authors provide useful data on the time responses. However, looking at those figures, it is puzzling why certain vessels were selected as responding as there seems almost no change after stimulation. In addition, some of the responses seem to actually start several tens of seconds before the actual stimulus (particularly in A).

Only some traces in C and D (dark blue) seem to be actually responding vessels.

This is not discussed and unclear.

Supplementary Figure 2 displays the time courses of vessel calibre for all vessels in the FOV, not just those deemed responders.

The aforementioned effects are due to the loess smoothing filter having been applied to the time courses for the preliminary response, which has been rectified in the updated figures. In particular, Supplementary Figure 2 has been updated with separate loess smoothing before and after photostimulation. The (pre-stimulation) effect is gone once the loess smoothing has been separated.

- R Point 7: As indicated before and in agreement with the alternative reviewer, the quality of the results in 3d is difficult to judge. No 2d sections that compare 'ground truth' with inferred results are shown in the current manuscript which would enable a much better judgment. The provided video is still 3d and not a video going through 2d slices. Also, in the video the overlap of vasculature and raw data seems to be very good and near 100%, why is the dice measure reported earlier so low ? Is this a particularly good example ?

Some examples, indicating where the pipeline fails (and why) would be helpful to see, to judge its performance better (ideally in 2d slices).

As discussed in the public comments, the 2D slices are now included in Suppl. Fig. 4 and suppl. Fig 13 to facilitate visual assessment. The vessels are long and thin so that slight dilations or constrictions impact the Dice scores without being easily visualizable.

- Author response images 6 and 7. From the presented data the constrictions measured in the smaller vessels may be a result (at least partly) of noise. This seems to be particularly the case in Author response image 7 left top and bottom for example. It would be helpful to show the actual estimates of the vessels radii overlaid in the (raw) images. In some of the pictures the noise level seems to reach higher values than the 10-20% of noise used in the tests by the authors in the revision.

The vessel radii are estimated as averages across all vertices of the individual vessels: it is thus not possible to overlay them meaningfully in 2D slices: in Figure 2B, we do show a rendering of sample vessel-wise radii estimates.

- "We tested the centerline detection in Python, scipy (1.9.3) and Matlab. We found that the Matlab implementation performed better due to its inclusion of a branch length parameter for the identification of terminal branches, which greatly reduced the number of false branches; the Python implementation does not include this feature (in any version) and its output had many more such "hair" artifacts. Clearmap skeletonization uses an algorithm by Palagyi & Kuba(1999) to thin segmentation masks, which does not include hair removal. Vesselvio uses a parallelized version of the scipy implementation of Lee et al. (1994) algorithm which does not do hair removal based on a terminal branch length filter; instead, Vesselvio performs a threshold-based hair removal that is frequently overly aggressive (it removes true positive vessel branches), as highlighted by the authors."

This statement is wrong. The removal of small branches in skeletons is algorithmically independent of the skeletonization algorithm itself. The authors cite a reference concerned with the algorithm they are currently employing for the skeletonization. Careful assessment of that reference shows that this algorithm removes small length branches after skeletonization is performed. This feature is available in open-source packages as well, or could be easily implemented.

We appreciate that skeletonization is distinct from hair removal and have reworded this paragraph for clarity. We are currently working with SciPy developers to implement hair removal in their image processing pipeline so as to render our pipeline fully open-source.

The removal of hairs after skeletonization with length based thresholding leads to the possibility of removing parts of centerlines in the main part of vessels after branch points with hairs. The Matlab implementation does not do this and leaves the main branches intact.

This text has been updated to:

“Hair” segments shorter than 20 μm and terminal on one end were iteratively removed, starting with the shortest hairs and merging the longest hairs at junctions with 2 terminal branches with the main vessel branch to reduce false positive vascular branches and minimize the amount of centerlines removed. This iterative hair removal functionality of the skeletonization algorithm is currently unavailable in python, but is available in Matlab [9].

- "On the reviewer's comment, we did try inputting normalized images into Ilastik, but this did not improve its results." This is surprising. Reasonable standard preprocessing (e.g. background removal, equalization, and vessel enhancement) would probably restore most of illastik's performance in the indicated panel.

While the improvement may be present in a particular set of images, the generalizability of such improvement to other patches is often poor in our experience, as reflected by aforementioned results and the widespread uptake of DL approaches to image segmentation. The in vivo datasets also contain artifacts arising from eg. bleeding into the FOV that ilastik is highly sensitive to. This is an example of noise that is not easily removed by standard preprocessing.

- "Typical pre-processing/standard computer vision techniques with parameter tuning do not generalize on out-of-distribution data with different image characteristics, motivating the shift to DL-based approaches."

I disagree with this statement. DL approaches can generalize typically when trained with sufficient amount of diverse data. However, DL approaches can also fail with new out of distribution data. In that situation they only be 'rescued' via new time intensive data generation and retraining. Simple standard image pre-processing steps (e.g. to remove background or boost vessel structures) have well defined parameter that can be easily adapted to new out of distribution data as clear interpretations are available. The time to adapt those parameters is typically much smaller than retraining of DL frameworks.

We find that the standard image processing approaches with parameter tuning work robustly only if fine-tuned on individual images; i.e., the fine-tuning does not generalize across datasets. This approach thus does not scale to experiments yielding large image sizes/having high throughput experiments. While DL models may not generalize to out-of-distribution data, fine-tuning DL models with a small subset of labels generally produce superior models to parameter tuning that can be applied to entire studies. Moreover, DL fine-tuning is typically an efficient process due to very limited labelling and training time required.

- It is still unclear how the authors pipeline performs compared with other (open source or commercially) available pipelines. As indicated before, comparing to illastik, particularly when feeding non preprocessed data, does not seem to be a particularly high bar.

This question has also been raised by the other reviewer who asked to compare to commercially available pipelines.

This question was not answered by the authors, and instead the authors reply by claiming to provide an open source pipeline. In fact, the use of matlab in their pipeline does not make it fully open-source either. Moreover, as mentioned before, open-source pipelines for comparisons do exists.

As discussed above, the manuscript now includes comparisons to Imaris for segmentation and Vesselvio for graph extraction. The pipeline is on github.

-"We agree with the review that this question is interesting; however, it is not addressable using present data: activated neuronal firing will have effects on their postsynaptic neighbors, yet we have no means of measuring the spread of activation using the current experimental model."

Distances to the closest neuron in the manuscript are measured without checking if it's active. Thus, distances to the first set of n neurons could be measured in the same way, ignoring activation effects.

Shorter distances to an entire ensemble of neurons would still be (more) informative of metabolic demands.

This could indeed be done within the existing framework. The connected-components-3d can be used to extract individual occurrences of neurons in the FOV from the neuron segmentation mask. Each neuron could then have its distance calculated to each point on the vessel centerlines.

- model architecture:

It is unclear from the description if any positional encoding was used for the image patches.

It is unclear if the architecture / pipeline can handle any volume sizes or is trained on a fixed volume shapes? In the latter case how is the pipeline applied?

The model includes positional encoding, as described in Hatamizadeh et al. 2021.

The model can be applied to images of any size, as demonstrated on larger images in Supplementary Figure 9 and on smaller images in Supplementary Figure 2. The pipeline is applied in the same way. It will read in the size of an input image and output an image of the same size.

- transformer models often show better results when using a learning rate scheduler that adjust the learning rate (up and down ramps typically). Did the authors test such approaches?

We did not use a learning rate scheduler, as we found we were getting good results without using one.

- formula (4): The 95% percentile of two numbers is the max, and thus (5) is certainly not what the HD95 metric is. The formula is simply wrong as displayed.

Thank you. The formula has been updated.

- formula (5): formula 5 is certainly wrong: n_X, n_y are either integer numbers as indicated by the sum indices or sets when used in the distances, but can't be both at the same time.

Thank you for your comment. The Formula has been updated.

- The statement:

"this functionality of the skeletonization algorithm is currently unavailable in any python implementation, but is available in Matlab [56]."

is not correct (see reply above)

Please see the response above. This text has been updated to:

“Hair” segments shorter than 20 μm and terminal on one end were iteratively removed, starting with the shortest hairs and merging the longest hairs at junctions with 2 terminal branches with the main vessel branch to reduce false positive vascular branches and minimize the amount of centerlines removed. This iterative hair removal functionality of the skeletonization algorithm is currently unavailable in Python, but is available in Matlab [9].

- the centerline extraction is performed after taking the union of smoothed masks. The union operation can induce novel 'irregular' boundaries that degrade skeletonization performance. I would expect to apply smoothing after the union?

Indeed the images were smoothed via dilation after taking the union, as described in the previous set of responses to private comments.

- "The radius estimate defined the size of the Gaussian kernel that was convolved with the image to smooth the vessel: smaller vessels were thus convolved with narrower kernels."

It's unclear what image were filtered ?

We have updated this text for clarity:

The radius estimate defined the size of the Gaussian kernel that was convolved with the 2D image slice to smooth the vessel: smaller vessels were thus convolved with narrower kernels.

- Was deconvolution on the raw images applied or after Gaussian filtering ?

The deconvolution was applied before Gaussian filtering.

- ",we extracted image intensities in the orthogonal plane from the deconvolved raw registered image. A 2D Gaussian kernel with sigma equal to 80% of the estimated vessel-wise radius was used to low-pass filter the extracted orthogonal plane image and find the local signal intensity maximum searching, in 2D, from the center of the image to the radius of 10 pixels from the center."

Would it not be better to filter the 3d image before extracting a 2d plane and filter then ?

That could be done, but would incur a significant computational speed penalty. 2D convolutions are faster, and produced excellent accuracy when estimating radii in our bead experiment.

What algorithm was used to obtain the 2d images.

The 2d images were obtained using scipy.ndimage.map_coordinates.

- Figure 2: H is this the filtered image or the raw data ?

Panel H is raw data.

- It would be good to see a few examples of the raw data overlaid with the radial estimates to evaluate the approach (beyond the example in K).

Additional examples are shown in Figure 5.

- Figure 2 K: Why are boundary points greater than 2 standard deviations away from the mean excluded ?

They are excluded to account for irregularities as vessels approach junctions [10], [11] REF.

- Figure 2 L: what exactly is plotted here ? What are vertex wise changes, is that the difference between the minimum and maximum of all the detected radii for a single vertex? Why do some vessels (red) show high values consistently throughout the vessel ?

Figure 2L displays change in the radius of vertices - in this FOV- following photostimulation in relation to baseline.

- Assortativity: to calculate the assortativity, are radius changes binned in any form to account for the fact that otherwise, $e_{xy}$ and related measures will be likely be based on single data points?

Assortativity is not calculated from single data points. It can be calculated by either binning into categories or computing it on scalars i.e. average radius across a vessel segment:

See here for info on calculating assortativity from binned categories (ie classifying a vessel as a constrictor, dilator or non-responder):

https://networkx.org/documentation/stable/reference/algorithms/generated/networkx.algorithms.assortativity.attribute_assortativity_coefficient.html#networkx.algorithms.assortativity.attribute_assortativity_coefficient

And see here for calculating assortativity from scalar values:

https://networkx.org/documentation/stable/reference/algorithms/generated/networkx.algorithms.assortativity.numeric_assortativity_coefficient.html#networkx.algorithms.assortativity.numeric_assortativity_coefficient

We calculated the assortativity using scalar values.

In both cases, one uses all nodes and calculates the correlation between each node and its neighbours with an attribute that can be binned or is a scalar. Binning the value on a given node would not affect the number of nodes in a graph.

- "Ilastik tended to over-segment vessels, i.e. the model returned numerous false positives, having a high recall (0.89{plus minus}0.19) but low precision (0.37{plus minus}0.33) (Figure 3, Supplementary Table 3)."

As indicated before, and looking at Figure 4, over segmentation seems due to too high background. A suggested preprocessing step on the raw images to remove background could have avoided this.

The images were normalized in preprocessing.

- Figure 4: The 3d panels are not much easier to read in the revised version. As suggested by other reviewers, 2d sections indicating the differences and errors would be much more helpful to judge the pipelines quality more appropriately.

As discussed above, 2D sections are now available in a supplementary figure.

- Figure 3: What would be the dice score (and other measures) between two ground truths extracted by two annotations by two humans (assisted e.g. by illastik).

Two additional human rates annotated images. We observed a ICC of 0.73 across a total of three raters on the three images.

- Figure 5: The authors only provide the absolute value of SU for the sigma noise levels. This only has some meaning when compared to the mean or median SU of the images. In the text the maximal intensity of 1023 SU is mentioned, but what are those values in images with weaker / smaller vessels (as provided in the constriction examples in the revision)/

I am unclear why this validation figure should be part of the main manuscript while generalization performance is left out.

The manuscript has been updated with the mean SNR value of 5.05 ± 0.15 to provide context for the quality of our images.

Bibliography

(1) J. R. Bumgarner and R. J. Nelson, “Open-source analysis and visualization of segmented vasculature datasets with VesselVio,” Cell Rep. Methods, vol. 2, no. 4, Apr. 2022, doi: 10.1016/j.crmeth.2022.100189.

(2) G. Tetteh et al., “DeepVesselNet: Vessel Segmentation, Centerline Prediction, and Bifurcation Detection in 3-D Angiographic Volumes,” Front. Neurosci., vol. 14, Dec. 2020, doi: 10.3389/fnins.2020.592352.

(3) N. Holroyd, Z. Li, C. Walsh, E. Brown, R. Shipley, and S. Walker-Samuel, “tUbe net: a generalisable deep learning tool for 3D vessel segmentation,” Jul. 24, 2023, bioRxiv. doi: 10.1101/2023.07.24.550334.

(4) W. Tahir et al., “Anatomical Modeling of Brain Vasculature in Two-Photon Microscopy by Generalizable Deep Learning,” BME Front., vol. 2020, p. 8620932, Dec. 2020, doi: 10.34133/2020/8620932.

(5) R. Damseh, P. Delafontaine-Martel, P. Pouliot, F. Cheriet, and F. Lesage, “Laplacian Flow Dynamics on Geometric Graphs for Anatomical Modeling of Cerebrovascular Networks,” ArXiv191210003 Cs Eess Q-Bio, Dec. 2019, Accessed: Dec. 09, 2020. [Online]. Available: http://arxiv.org/abs/1912.10003

(6) T. Jerman, F. Pernuš, B. Likar, and Ž. Špiclin, “Enhancement of Vascular Structures in 3D and 2D Angiographic Images,” IEEE Trans. Med. Imaging, vol. 35, no. 9, pp. 2107–2118, Sep. 2016, doi: 10.1109/TMI.2016.2550102.

(7) T. B. Smith and N. Smith, “Agreement and reliability statistics for shapes,” PLOS ONE, vol. 13, no. 8, p. e0202087, Aug. 2018, doi: 10.1371/journal.pone.0202087.

(8) J. R. Mester et al., “In vivo neurovascular response to focused photoactivation of Channelrhodopsin-2,” NeuroImage, vol. 192, pp. 135–144, May 2019, doi: 10.1016/j.neuroimage.2019.01.036.

(9) T. C. Lee, R. L. Kashyap, and C. N. Chu, “Building Skeleton Models via 3-D Medial Surface Axis Thinning Algorithms,” CVGIP Graph. Models Image Process., vol. 56, no. 6, pp. 462–478, Nov. 1994, doi: 10.1006/cgip.1994.1042.

(10) M. Y. Rennie et al., “Vessel tortuousity and reduced vascularization in the fetoplacental arterial tree after maternal exposure to polycyclic aromatic hydrocarbons,” Am. J. Physiol.-Heart Circ. Physiol., vol. 300, no. 2, pp. H675–H684, Feb. 2011, doi: 10.1152/ajpheart.00510.2010.

(11) J. Steinman, M. M. Koletar, B. Stefanovic, and J. G. Sled, “3D morphological analysis of the mouse cerebral vasculature: Comparison of in vivo and ex vivo methods,” PLOS ONE, vol. 12, no. 10, p. e0186676, Oct. 2017, doi: 10.1371/journal.pone.0186676.

Changes from the earlier version to the revised version

Abstract

Abstract

Functional hyperaemia is a well-established hallmark of healthy brain function, whereby local brain blood flow adjusts in response to a change in the activity of the surrounding neurons. Although functional hyperemia has been extensively studied at the level of both tissue and individual vessels, vascular network-level coordination remains largely unknown. To bridge this gap, we developed a deep learning-based computational pipeline that uses two-photon fluorescence microscopy images of cerebral microcirculation to enable automated reconstruction and quantification of the geometric changes across the microvascular network, comprising hundreds of interconnected blood vessels, pre and post-activation of the neighbouring neurons. The pipeline’s utility was demonstrated in the Thy1-ChR2 optogenetic mouse model, where we observed network-wide vessel radius changes to depend on the photostimulation intensity, with both dilations and constrictions occurring across the cortical depth, at an average of 16.1±14.3 μm (mean±stddev) away from the most proximal neuron for dilations; and at 21.9±14.6 μm away for constrictions. We observed a significant heterogeneity of the vascular radius changes within vessels, with radius adjustment varying by an average of 24 ± 28% of the resting diameter, likely reflecting the heterogeneity of the distribution of contractile cells on the vessel walls. A graph theory-based network analysis revealed that the assortativity of adjacent blood vessel responses rose by 152 ± 65% at 4.3 mW/mm2 of blue photostimulation vs. the control, with a 4% median increase in the efficiency of the capillary networks during this level of blue photostimulation in relation to the baseline. Interrogating individual vessels is thus not sufficient to predict how the blood flow is modulated in the network. Our computational pipeline, to be made openly available, enables tracking of the microvascular network geometry over time, relating caliber adjustments to vessel wall-associated cells’ state, and mapping network-level flow distribution impairments in experimental models of disease.

Introduction

To support healthy brain functioning, the cerebrovascular network undergoes continual adjustments in vessel calibers [1], [2], [3]. Neurovascular coupling refers to the change in blood flow following changes in the level of neuronal activity: under physiological conditions, a generous buffer of nutrients is granted to activated parenchyma via the capillary network [1], [4]. This buffer is maintained through finely tuned regulation of flow through changes in the vessel caliber, mediated via contractile cells in the vessel walls. In the absence of such tuning, pockets of tissue could experience inadequate access to metabolites [5]. Alterations in smooth muscle cells, pericytes, and astrocytes may lead to compromises in vessels’ dilatory capacity and thus deficits in neurovascular coupling [6], [7], [8], [9]. In various brain pathologies, including Alzheimer’s disease, stroke, and trauma, regional blood flow regulation gets impaired through vessel loss and/or dysfunction of the vessels’ dilatory capacity, resulting in regions of ischemia/hypoxia [10], [11]. Previous studies have examined either individual vessels or the tissue level responses, with little attention having been paid to the vascular network, though network dysfunction frequently is associated with accelerated disease progression and long-term symptomatology [3], [12], [13], [14], [15], [16], [17], [18], [19], [20].

While there is copious data on the functioning of individual vessels, interrogation of the microvascular network remains a challenge, in terms of both data acquisition and analysis [6], [7], [21], [22], [23]. To date, studies on the brain vasculature have been done by sparsely imaging individual blood vessels at the cellular scale [6], [7], [21], [24], [25], [26], [27], [28], [29], [30], [31], [32], thereby severely undersampling the microvascular network; or by evaluating the averaged flow over many vessels at the mesoscopic scale, thus failing to discern the flow through individual vessels. A critical gap in the field is the characterization of flow across hundreds of individual vessels, while imaging the network structure that links them together to determine how the vascular response is coordinated across the network. This gap is particularly significant as studies investigating blood flow across several vessels at a time (imaged by varying the line acquisition pattern [6], [24]) have shown highly heterogeneous responses among capillaries. Neuronal function impairments arise wherever local metabolite supply becomes inadequate, notwithstanding the physiological level of flow across the network as a whole, making mapping of vessel changes across the network of particular importance.

To address the limitations of previous work, we developed a novel deep learning (DL)-based pipeline for mapping changes to the geometry of the brain vascular network following neuronal activation, from a time series of volumetric two-photon fluorescence microscopy (2PFM) data. Neuronal activation was elicited by photostimulation of pyramidal neurons expressing Channelrhodopsin-2 (ChR2) [33] in the Thy1-ChR2-YFP mouse model [34]. Our DL pipeline enabled automatic and accurate segmentation, registration and network analysis of large 2PFM datasets across time. We applied our pipeline in a dataset of 17 Thy1-ChR2-YFP mice to map photostimulation-induced changes across the microvascular network - at the level of individual vessels and at the level of vertices spaced every micrometre along the vessels – in relation to the distance to the closest pyramidal neurons expressing the optogenetic actuator and across the cortical depth. Our findings demonstrate the utility of our pipeline for studying in situ microvascular morphology and function to address various neuroscientific hypotheses.

Methods

Cohort

Animals

All experimental procedures in this study followed the ARRIVE 2.0 guidelines [35]. They were approved by the Animal Care Committee of the Sunnybrook Research Institute, which adheres to the policies and guidelines of the Canadian Council on Animal Care and meets all the requirements of the Provincial Statute of Ontario, Animals for Research Act, and the Canadian Federal Health of Animals Act. We used 15 male and 13 female Thy1-ChR2-YFP mice (#007612, line 18, Jackson Laboratory) at 6-12 months of age (283.9 ± 67.0 days), weighing 28.8±7.1g on the day of imaging. The mice were bred in-house and housed under a 12-hr light/dark cycle [34]. In this mouse line, Channelrhodopsin-2 is expressed preponderantly by pyramidal neurons with soma in cortical layer 5 and dendrites projecting to the cortical surface, enabling depolarization of their cellular membranes and action potential generation upon blue light illumination [30], [34]. An attrition schematic is provided in Supplementary Figure 1.

Surgical Preparation

On the day of imaging, mice were induced with 5% isoflurane, transferred to a rectal probe feedback-regulated heating pad (CWE Inc, Ardmore PA) to maintain a temperature of 37 °C, and maintained under 2% isoflurane. A subcutaneous injection of 1 mL of lactated Ringer’s solution was administered at the start of surgery. Throughout the surgical preparation and imaging, we monitored breath rate, heart rate, arterial blood oxygen saturation, pulse distention, and breath distention via a pulse oximeter, with a probe mounted on the thigh (MouseOx, STARR Life Sciences) (Supplementary Table 1). For fine control over respiration, the mice were tracheostomized with an endotracheal tube (20 Ga catheter) and ventilated with a gas mixture of 20-30% O2 and medical air using a small animal ventilator (SAR 830/P, CWE Inc, Ardmore PA) set to 115-130 breaths per minute at an inspiratory/expiratory ratio of 1:3-4. Their heads were immobilized via ear bars during the placement of the cranial window and imaging. The tail vein was cannulated with a 26 Ga catheter to enable fluorophore and anesthetic delivery. A cranial window was implanted over the forelimb region of the primary somatosensory cortex (AP 0.25mm, ML 2.0 mm); the dura was excised, and 1% agarose was applied between the brain and the glass coverslip. A well was built using dental cement to allow water immersion of the objective. Texas Red dextran (70 kDa MW, Thermo Fisher Scientific Inc, Waltham MA) was diluted in PBS and injected through the tail vein catheter at a concentration of 33 mg/kg for a total injection volume between 0.1 and 0.2 mL. A line containing 0.01 g/ml alpha chloralose was connected to the tail vein catheter. The animal was then positioned underneath the microscope objective, and the structural 2PFM scan acquired, as detailed below. Seventeen mice (nine male and eight female) made it through the surgical preparation (cf. Supplementary Figure 1). Following the structural acquisition, the isoflurane was discontinued, and the continuous infusion of alpha chloralose commenced at 40 mg/kg/hr.

Imaging

Two-photon Fluorescence Imaging and Optogenetic Stimulation

Mice were imaged on an FVMPE-RS microscope (Olympus, Japan) using a 25x/1.05NA objective. An Insight tunable Ti:Saphire near-infrared laser (SpectraPhysics, USA) was used for 900 nm excitation of Texas Red-labelled vasculature and YFP-labeled ChR2-expressing pyramidal neurons. Two visible light stimulation lasers were used to excite ChR2 at either 458 nm or 552 nm (control). The optical setup is shown in Figure 1. The imaging field-of-view was selected to include at least one penetrating artery and vein yet avoid major pial vessels and thus signal loss in the underlying cortex. Structural scans were acquired under 2% isoflurane using a Galvano scanner, with 2x averaging, a z-step of 0.99 μm, and a nominal lateral resolution of 0.99 μm. Following structural scanning, mice were allowed to rest for 5 minutes to let blood flow equilibrate from brain exposure to room light. The five-minute resting period was observed due to prior reports of red blood cell velocities following optogenetic ChR2 stimulation in the same mouse strain remaining altered for a minute following photostimulation and due to potential pericyte-mediated responses via cytoskeleton reorganization occurring over a minute [6], [30].

Photostimulation setup:

The excitation and stimulation light pass through a FV30-NDM690 dichroic mirror with two notch filters, at 458 nm and 552 nm, to excite TexasRed, EYFP and ChR2 within the mouse. The emitted light passes through the objective, is reflected off the FV30-NDM690 dichroic mirror, and passes through a 650 nm barrier filter before reaching a 570 nm long pass filter (LPF) separating emitted light from EYFP and TexasRed, which respectively pass through 495-540 nm and 575-630 nm barrier filters to be collected via GaAsP detectors.

Functional scans were acquired over baseline and post-stimulus periods, flanking a period of blue light illumination. The haemodynamic response to optogenetic stimulation in the Thy1-ChR2-EYFP mouse model was previously observed to last about 45 seconds, and observed to last for minutes in our experimental model (cf. Supplementary Figure 2) [30] These vascular radius time courses were similar to data collected in other optogenetic models of mural cell stimulation where arteries were observed to dilate for approximately 60 and 20 seconds respectively, and capillaries were observed to dilate on the scale of minutes [6], [23]. We thus constrained our volumetric acquisitions to 45 seconds. Each volumetric scan lasted for 42.98 seconds and used the resonant scanner with 5x frame averaging (resulting in a signal-to-noise ratio on the vasculature channel of 6.9±1.6), a nominal lateral resolution of 0.99 μm, a z-step of 2.64 μm, and traversed 250.8 μm of cortical depth. This z-step was chosen because larger z-steps led to discontinuities in the capillary bed segmentation. We imaged the brain in sections from the cortical surface to a depth of 250 μm and from 250 μm to 500 μm of cortical depth. The acquisitions were split into two slabs to maximize the volume coverage without creating discontinuities in the capillary segmentation masks while maintaining a scan duration of under 45 seconds. Stacks were acquired either from the cortical surface down or from the bottom of the stacks toward the cortical surface on different paired acquisitions, so that the post-stimulus delay of image acquisition at different cortical depths was not constant. A 239-μm diameter circular photostimulation region was positioned 250 μm beneath the cortical surface. The photostimulation light was raster scanned over this ROI for 5 seconds between imaging scans, with a pixel dwell time of 4 μs for each pixel of the stimulation ROI. Each pixel in the photostimulation ROI was thus excited every 501.55 milliseconds, i.e. at about 2Hz. We used two laser powers for 458-nm photostimulation: 1.1 mW/mm2 and 4.3 mW/mm2. The 1.1 mW/mm2 is just above the threshold for ChR2 photoactivation and is expected to elicit a small degree of neuronal activity, whereas the 4.3 mW/mm2 photoactivation elicits more neuronal activity while still being below ChR2 saturation level [33], [36]. Low powers were used to activate ChR2-expressing neurons to minimize tissue heating [37]. We also performed control experiments via 552-nm photostimulation (where ChR2 excitation falls to 4% of its peak) at 4.3 mW/mm2 [38]. Mice rested for approximately 5-7 minutes between scan pairs for the vascular tone to return to its resting state. Photostimulation parameters were presented in randomized order for each mouse.

Segmentation and Graph Extraction

Ground truth generation

To generate ground truths, we used ilastik’s pixel classification workflow, which relies on random forests, to annotate blood vessels and pyramidal excitatory neurons soma’s in 42 volumes from 25 Thy1-ChR2-EYFP mice for training our DL model (24 images from 16 Thy1-ChR2-EYFP mice from other studies were included in the training, validation, and test cohorts to increase our dataset size). Images were semi-automatically segmented in groups of up to 4 each, with a size of 507×507×250 μm each, due to ilastik’s inability to handle large amounts of annotations over more than a few volumes. The rater labelled targets (neurons vs. blood vessels), corrected mistakes, and classified the pixels where the output was uncertain (as shown using ilastik’s uncertainty guidance feature). During the manual annotation, feature selection was repeatedly optimized using ilastik’s suggest feature function (wrapper method), leading to different features being used for the random forest model for each set of images. Small (under 50 pixels) isolated vascular components were removed with connected component analysis in Python using scikit-image and connected-components-3d. Testing data was withheld until the final model was selected based on the models’ performance evaluation on the validation data set.

Data preprocessing and augmentation

We used the MONAI python package for data preprocessing, augmentation, model training, and prediction [39]. All images and ground truth segmentation masks were up-sampled to an isotropic voxel size using bilinear and nearest neighbour interpolation. Raw 10-bit image intensities were normalized to range from 0 to 1.0. Each volume had eight 128×128×128 pixel patches randomly cropped out for training. Data augmentation, transformations and parameters are listed in Supplementary Table 2. Spatial transformations were selected to expand data variety via cropping, rotations, and mirroring hence exposing the network to images that would be acquired on different positioning of the animal under the microscope. Zooming and deformation transformations were included to expose the network to small changes in the size and morphology of the vasculature. Intensity and Gaussian transformations exposed the network to signal intensity and contrast variations. Dropping pixels was included as the resonant scanner acquisition occasionally yields images with some relatively low signal pixels.

Model Architecture

Earlier version: We trained a state-of-the-art 3D vision transformer (UNETR) model and the U-Net model for baseline comparison, as implemented in PyTorch using the MONAI library [39]. UNETR is a U-Net style architecture with a transformer-based multi-attention head encoder and a CNN decoder [40], [41]. The encoder takes an input image (or patch from each batch) and breaks it down into a sequence of non-overlapping patches, each of size 16×16×16 pixels, which are weighted differently to account for variation in signal intensities and patterns within an image. The sequence of patches is then passed through a multi-head self-attention and multilayer perceptron encoder to capture self-attention between different pixels of the patch to encode long-range relationships between patches. The encoder is then connected to a CNN decoder with skipped connections to the encoder to map features back onto the original image at multiple spatial scales [40]. The model was trained on 2-channel images (EYFP-expressing neurons and Texas Red-labelled vasculature) as inputs to segment neurons and vasculature. We chose the following parameters for the UNETR architecture: a 12 multi-attention head encoder, 16 convolutional features in the first layer of the encoder, a hidden layer size of 768, a multilayer perceptron size of 3072, and Monte Carlo dropout, where on each run of the model 10% of weights were zeroed [40], [41], [42].

Revised version: We trained a state-of-the-art 3D vision transformer (UNETR) model and the U-Net model for baseline comparison, as implemented in PyTorch using the MONAI library [39]. UNETR is a U-Net style architecture with a transformer-based multi-attention head encoder and a CNN decoder [40], [41]. The encoder takes an input image (or patch from each batch) and breaks it down into a sequence of non-overlapping patches, each of size 16×16×16 pixels, which are weighted differently to account for variation in signal intensities and patterns within an image.

Added in the revised version: The sequence of patches is then passed through a multi-head self-attention and multilayer perceptron encoder to capture self-attention between different pixels of the patch to encode long-range relationships between patches. The encoder is then connected to a CNN decoder with skipped connections to the encoder to map features back onto the original image at multiple spatial scales [40]. The model was trained on 2-channel images (EYFP-expressing neurons and Texas Red-labelled vasculature) as inputs to segment neurons and vasculature. We chose the following parameters for the UNETR architecture: a 12 multi-attention head encoder, 16 convolutional features in the first layer of the encoder, a hidden layer size of 768, a multilayer perceptron size of 3072, and Monte Carlo dropout, where on each run of the model 10% of weights were zeroed [40], [41], [42].

The U-net model used is based on the original 2D U-net proposed by Ronneberger et al. in 2015 and extended into 3D via 3D convolutional operations and features residual blocks, parametric rectifying linear units, and instance normalization as described by Kerfoot et al. [41], [43], [44], [45], [46], [47]. The residual blocks were implemented for better gradient flow during the network training. At the same time, the parametric ReLU activation functions were used to improve the ability of the model to adjust its weights during training. Instance normalization was implemented to reduce contrast differences between images fed into the network to improve model robustness. The model was implemented with five layers featuring 16, 32, 64, 128, and 256 channels and dropout. The U-net model was trained using the same data augmentation employed during the UNETR model training.

Model Hyperparameter Optimization, Training, and Prediction

We used a grid search to determine optimal hyperparameters for the UNETR and UNet models within the following parameter space [43], [48]: loss functions (Dice Loss, Dice + Cross Entropy Loss, Dice + Focal Loss, or Tversky Loss), dropout rates (0.1, 0.2, or 0.3), learning rates (5e-3, 1e-3, 5e-4, 1e-4, or 1e-5), and the number of residual units for the Unet model (2 or 3) [49], [50]. We utilized the Adam optimizer [51], and the best model during hyperparameter optimization was selected based on the validation Dice similarity coefficient (DSC) and trained for a maximum of 2400 epochs. The Dice score was the principal evaluation metric, to maximize the overlap between ground truth and prediction masks. Precision and recall were used as secondary metrics to achieve balanced precision and recall where over-segmented results were produced. For early stopping, the epoch with the best performance of the DSC on the validation dataset was selected. The final model that had the best Dice score during hyper-parameter optimization was trained with a Dice + Cross Entropy Loss function, a dropout rate of 0.1, 2 residual units, a learning rate of 1e-5, and a batch size of 1 image with 8 crops per image. Training and optimization were performed on the Narval cluster of Calcul Quebec and the Digital Research Alliance of Canada, with each node using 498 GB of RAM and 4 Nvidia A100 GPUs, each with 40GB HBM2 VRAM.

Ilastik utilized a random forest model from the Vigra library with the default parameters and 100 trees [52], [53]. We initially utilized all default 3D Colour/Intensity/Texture/Edge features during ilastik feature selection and added 2D Colour/Intensity/Texture/Edge features with a sigma of 20.0 pixels. These 2D features were added as a strong, largely planar artifact was observed towards the surface of the images in the neuron channel. We used the wrapper method for feature selection with a set size penalty of 0.10 to determine the optimal features from the starting set. The model was then trained using live updates on the Narval cluster of Calcul Quebec and the Digital Research Alliance of Canada, with each node using 249 GB of RAM and 2 x AMD Rome 7532 CPUs with 64 cores.

For the ilastik model, a subset of 3 training images from 3 mice was used as ilastik’s random forest was unable to train a model with more data even on compute nodes with large memory and CPU resources. The model could never complete training within the time limits of the resource allocations. This inability to complete model optimization was expected as ilastik’s documentation does not recommend training random forest models with full annotation datasets, and increasing the number of annotations does not necessarily lead to better predictions for the pixel classification workflow [52]. The ilastik random forest model was trained with whole images rather than with patches.

Model Comparison

To compare models, we employed several metrics evaluating the similarity between the ground truth and prediction: the Dice Similarity Coefficient (DSC), Precision, and Recall, as well as surface-based metrics: 95% Hausdorff distance (HD95) and mean surface distance. We specifically focused on mean surface distance and HD95 distance since the centerline extraction used in the downstream analysis was highly sensitive to surface irregularities. The model with the lowest standard deviation on the mean surface distance was to be selected absent statistically significant differences in the average mean surface distance between different models.

Removed in the revised version: Consistency in surface placement was key for extracting good centerlines and graphs from the segmentation masks. The metrics utilized are defined as follows:

Earlier version: The DSC measures the overlap between the prediction and ground truth with a value of 1 corresponding to complete overlap, X represents the ground truth and Y represents the predicted value. Precision measures the rate of correctly returned predicted values, whereas recall accesses the rate of return of targeted values. For the Hausdorff 95% distance (HD95), the distance d is the infimum of the Euclidean distance between points x and y (from the surface of segmentation masks X and Y) to segmentation masks Y and X, respectively. HD95 computes the maximum of the 95th percentile of the supremums (sup) of these minimum surface distances (d(x,Y) and d(X,y)) between the boundaries of the ground truth and the predicted segmentation mask. The mean surface distance measures the average minimum distance between the boundary of the ground truth and the boundary of the predicted segmentation mask. In Equation (5), nx and ny are the sets of boundary points of X and Y, respectively. The infimum of the minimum Euclidean distance, d, from boundary points p from nx, and q from ny to sets nx and ny are then computed and averaged to compute the Mean Surface Distance. Together, the HD95 and Mean Surface Distance assess model performance at the boundaries of the segmentation masks.

Revised version: Consistency in surface placement was key for extracting good centerlines and graphs from the segmentation masks. The metrics utilized are defined as follows: The DSC measures the overlap between the prediction and ground truth with a value of 1 corresponding to complete overlap, X represents the ground truth and Y represents the predicted value. Precision measures the rate of correctly returned predicted values, whereas recall accesses the rate of return of targeted values. For the Hausdorff 95% distance (HD95), the distance d is the infimum of the Euclidean distance between points x and y (from the surface of segmentation masks X and Y) to segmentation masks Y and X, respectively. HD95 computes the maximum of the 95th percentile of the supremums (sup) of these minimum surface distances (d(x,Y) and d(X,y)) between the boundaries of the ground truth and the predicted segmentation mask. The mean surface distance measures the average minimum distance between the boundary of the ground truth and the boundary of the predicted segmentation mask. In Equation (5), nx and ny are the sets of boundary points of X and Y, respectively, and Nx and Ny are total number of boundary points nx and ny. The infimum of the minimum Euclidean distance, d, from boundary points p from nx, and q from ny to sets nx and ny are then computed and averaged to compute the Mean Surface Distance. Together, the HD95 and Mean Surface Distance assess model performance at the boundaries of the segmentation masks.

Graph extraction of cerebrovascular networks

Graph extraction was performed in Python except for the centerline calculation, which was performed in MatLab (R2021a). Upsampled image acquisitions (0.99 μm isotropic voxel size) were registered with ANTsPy using the rigid registration method with a total sigma of 2 pixels (for smoothing within the registration function) and mean squared error as the similarity metric as input parameters of the registration function (Figure 2A,B). We selected the first baseline scan from each region of interest that was scanned from the bottom to the top to serve as a reference to which all other images from the same region were aligned. The calculated transforms were used to transform images from the same ROI to the space of the reference image using linear interpolation [54]. We then generated segmentation masks for each aligned image using our trained UNETR model with sliding window inference, and we retained Monte Carlo dropout during prediction to create an ensemble of 20 UNETR models, so as to assess the model uncertainty (Figure 2D) [55].

Computational analysis pipeline:

A. The stacks of 2PFM slices were registered using ANTS rigid registration and aligned to the reference time point. B. Images were upsampled using bicubic interpolation to an isotropic resolution of 0.99×0.99×0.99 μm. C. An ensemble of UNETR deep learning models with dropout generated segmentation masks at each time point, producing probability maps. D. The mean and standard deviation of the probability of each pixel being vasculature were computed and used to create binary vascular segmentation masks. E. The union over the vascular segmentation masks for all time points was computed, and background pixel clusters within vessel masks were removed. F. The vascular segmentation mask was thinned down to centerlines and rendered as a graph, where edges were vessel segments connecting branch points (nodes). This skeleton was overlaid on the vasculature channel from which the neuron channel was subtracted. G. The plane orthogonal to the tangent to the vessel’s travel direction was computed every micrometre along the centerline. H,I. 1D signal intensity profiles at each centerline vertex were computed in the orthogonal plane every 10°. J. The boundary for each profile was placed at the minimum of the signal gradient for that signal intensity profile. K. The raw intensity image with the detected boundary points, where outlier boundary points (in green) were defined as points over 2 standard deviations from the mean and excluded. L. Visualization of the changes in vertex-wise radii on a sample vascular network.

Earlier version: To extract the centerlines, we assumed that no background was present within the vessels, as regions of background within vessels disrupt the centerline extraction algorithm. Background labelled pixels surrounded by blood vessel segmentations were thus filled in using connected component analysis. Disconnected vascular components under 50 pixels were assumed to be noise and removed. Each aligned binary segmentation mask was dilated three times with a footprint defined by scikit-image’s disk function with a radius of 1 pixel. Next, we computed the union of all segmentation masks from the same ROI generated from registered images (Figure 2E). This common segmentation mask for all time points was eroded to a centerline via a medial axis transformation (bwskel function in Matlab (R2021a)). Critically, using the union of all time points minimized the sensitivity of centreline identification to red blood cell (RBC) stalls at individual time points. “Hair” segments shorter than 20 μm and terminal on one end were removed to reduce false positive vascular branches: this functionality of the skeletonization algorithm is currently unavailable in any python implementation, but is available in Matlab [56]. The vascular centerlines were next used to construct vascular graphs (with sknw) for each ROI, rendering coordinates of vertices along the vessel’s centreline as edges, and branch points as nodes [57].

Revised version: To extract the centerlines, we assumed that no background was present within the vessels, as regions of background within vessels disrupt the centerline extraction algorithm. Background labelled pixels surrounded by blood vessel segmentations were thus filled in using connected component analysis. Disconnected vascular components under 50 pixels were assumed to be noise and removed. Each aligned binary segmentation mask was dilated three times with a footprint defined by scikit-image’s disk function with a radius of 1 pixel. Next, we computed the union of all segmentation masks from the same ROI generated from registered images (Figure 2E). This common segmentation mask for all time points was eroded to a centerline via a medial axis transformation (bwskel function in Matlab (R2021a)). Critically, using the union of all time points minimized the sensitivity of centreline identification to red blood cell (RBC) stalls at individual time points. “Hair” segments shorter than 20 μm and terminal on one end were iteratively removed, starting with the shortest hairs and merging the longest hairs at junctions with 2 terminal branches with the main vessel branch to reduce false positive vascular branches and minimize the amount of centerlines removed. This iterative hair removal functionality of the skeletonization algorithm is currently unavailable in python, but is available in Matlab [56]. The vascular centerlines were next used to construct vascular graphs (with sknw) for each ROI, rendering coordinates of vertices along the vessel’s centreline as edges, and branch points as nodes [57].

Earlier version: Based on the vascular graphs, we computed vascular radius estimates at each vertex at each timepoint and then calculated the vertex’s distance to the nearest YFP-expressing pyramidal neuron (as measured by the distance transformation at the vertex to the nearest labelled neuron, with the radius of the vessel subtracted off). We employed a two-stage approach to estimate the radius for each point on the centerline. First, using the binary segmentation mask, we calculated the distance from every vessel pixel in the mask to the nearest background pixel. These values were averaged for each vessel to estimate its radius. The radius estimate defined the size of the Gaussian kernel that was convolved with the image to smooth the vessel: smaller vessels were thus convolved with narrower kernels.

Revised version: Based on the vascular graphs, we computed vascular radius estimates at each vertex at each timepoint and then calculated the vertex’s distance to the nearest YFP-expressing pyramidal neuron (as measured by the distance transformation at the vertex to the nearest labelled neuron, with the radius of the vessel subtracted off). We employed a two-stage approach to estimate the radius for each point on the centerline. First, using the binary segmentation mask, we calculated the distance from every vessel pixel in the mask to the nearest background pixel. These values were averaged for each vessel to estimate its radius. The radius estimate defined the size of the Gaussian kernel that was convolved with the 2d image slice to smooth the vessel: smaller vessels were thus convolved with narrower kernels.

In the second stage, the registered raw image was deconvolved with the point spread function of the microscope, as measured via FluoSpheres carboxylate-modified Microspheres (Cat# F8803, Thermo Fisher Scientific Inc, Waltham MA), using Richardson-Lucy deconvolution. To low pass filter the centreline in advance of computing the tangent vector at each vertex, the coordinates of the vertices along the centerline were smoothed using a Gaussian with a sigma of 3. Next, the tangent vector to the centerline was specified by calculating the gradient of the centerline path. Given this local tangent to the vessel’s direction of travel at a given vertex, we extracted image intensities in the orthogonal plane from the deconvolved raw registered image. A 2D Gaussian kernel with sigma equal to 80% of the estimated vessel-wise radius was used to low-pass filter the extracted orthogonal plane image and find the local signal intensity maximum searching, in 2D, from the center of the image to the radius of 10 pixels from the center. The orthogonal plane image was sampled every 10 degrees (as finer radial sampling did not improve the estimation cf. Supplementary Figure 3) along radial lines emanating from the local signal intensity maximum closest to the center of the image and 5x bicubic upsampled to extract thirty-six 1D signal intensity profiles at that vertex, as shown in Figure 2H. These line profiles were then convolved with a 1D Gaussian kernel with a sigma of 80% of the estimated radius of that vessel, and the gradient of each profile was calculated as shown in Figure 2I, J. The vessel boundary was then placed at the local minimum of the gradient of each profile. The mean and standard deviation of the boundary distances for the thirty-six 1D line profiles were calculated, and boundary points greater than 2 standard deviations away from the mean were excluded (Figure 2K). This radius estimation procedure was repeated for all vertices of all vessels.

In the last step, we computed the distance from each image voxel to the nearest YFP-expressing pyramidal neuron by computing the distance from every image pixel not belonging to the neuron mask to the closest YFP-expressing neuron’s soma boundary (based on the neuron’s binary segmentation mask). At the vessel vertices, these distances were adjusted by subtracting the vessel’s radius. We thereby captured the distance from the vessel surface to the closest YFP-expressing neuron.

Boundary detection validation

To assess the uncertainty in the radius estimates at each vertex, we simulated changes to the vascular diameter by “resizing” the extracted orthogonal plane image and adding Gaussian noise to it. These steps were undertaken to validate the ability of the boundary detection method to estimate diameter changes and evaluate the robustness of the estimates to increased amounts of noise. For the diameter change estimation, images were resampled from the orthogonal plane by a random factor, uniformly distributed in the range of 0.5x to 2x. Then, the aforementioned boundary detection methods were used to estimate vertex-wise calibre changes. The end goal of the assessment was to see how close the boundary detection method-based radius change matched the prescribed diameter change. In the second task, we added Gaussian noise with a sigma randomly chosen from a uniform distribution ranging from 0 to 500 SU to the orthogonal plane image. The noise was added to the image after deconvolution with the PSF and the extraction of the orthogonal plane but before Gaussian smoothing. We then reported on the percent change in the radius as a result of resizing or adding noise in relation to the baseline radius.

Added in the revised version: Second, our boundary detection algorithm was used to estimate the diameters of fluorescent beads of a known radius imaged under similar acquisition parameters. Polystyrene microspheres labelled with Flash Red (Bangs Laboratories, inc, CAT# FSFR007) with a nominal diameter of 7.32 μm and a specified range of 7.32 ± 0.27 μm as determined by the manufacturer using a Coulter counter were imaged on the same multiphoton fluorescence microscope set-up used in the experiment (identical light path, resonant scanner, objective, detector, excitation wavelength and nominal lateral and axial resolutions, with 5x averaging). The images of the beads had a higher SNR than our images of the vasculature, so Gaussian noise was added to the images to degrade the SNR to the same level of that of the blood vessels (SNR value of 5.05 ± 0.15). The images of the beads were segmented with a threshold, centroids calculated for individual spheres, and planes with a random normal vector extracted from each bead and used to estimate the diameter of the beads. The same smoothing and PSF deconvolution steps were applied in this task. We then reported the mean and standard deviation of the distribution of the diameter estimates. A variety of planes were used to estimate the diameters.

Vascular Network Analysis

Earlier version: Leveraging the graph representation of the vasculature in our pipeline, we next used graph theory to better understand the networks’ behaviour upon neuronal activation. We looked at morphometrics including vascular segment count density, vascular length density, and mean vessel length for comparison with other work. To demonstrate the benefits of extracting a graphical representation of the vasculature in situ, we looked at graph theory metrics including the assortativity of vascular radius changes and changes to the global efficiency of the capillary (below 5 μm in radius) network following optogenetic stimulation. The assortativity measures the tendency for one vessel to dilate or constrict by a similar amount as its neighbours. The assortativity (q) of radius changes in response to stimulation is defined as the Pearson correlation coefficient of these changes on connected vessels: where exy represents the fraction of vessels (edges) in the network that join together nodes with values x and y (i.e. radius changes with values x and y); ax and by are the percentages of edges connecting nodes with values x and y, and σa and σb are the standard deviations of ax and by [58]. The efficiency, in turn, captures how easily the graph can be traversed. For vascular graphs, efficiency can be conceptualized as the average of the inverse of the total resistive distance of the shortest paths of all combinations of vascular junctions, with the hydraulic resistivity serving as the distance between them. The resistivity (equation 7) is summed along the shortest paths, and the inverse of this sum is then averaged across all shortest paths to compute the efficiency (equation 8). Finally, photostimulation-induced change in the efficiency, from its baseline level, is reported.

Revised version: Leveraging the graph representation of the vasculature in our pipeline, we next used graph theory to better understand the networks’ behaviour upon neuronal activation. We looked at morphometrics including vascular segment count density, vascular length density, and mean vessel length for comparison with other work. To demonstrate the benefits of extracting a graphical representation of the vasculature in situ, we looked at graph theory metrics including the assortativity of vascular radius changes and changes to the global efficiency of the capillary (below 5 μm in radius) network following optogenetic stimulation. The assortativity measures the tendency for one vessel to dilate or constrict by a similar amount as its neighbours. The assortativity (q) of radius changes in response to stimulation is defined as the Pearson correlation coefficient of these changes on connected vessels: where exy represents the fraction of vessels (edges) in the network that join together nodes with values x and y (i.e. radius changes with values x and y); ax and by are the percentages of edges connecting nodes with values x and y, and σa and σb are the standard deviations of ax and by [58]. The efficiency, in turn, captures how easily the graph can be traversed. For vascular graphs, efficiency can be conceptualized as the average of the inverse of the total resistive distance of the shortest paths of all combinations of vascular junctions, with the hydraulic resistivity serving as the distance between them. The resistivity (equation 7) is summed along the shortest paths, and the inverse of this sum is then averaged across all shortest paths to compute the efficiency (equation 8). Finally, photostimulation-induced change in the efficiency, from its baseline level, is reported. In computing resistivity, we assumed a fluidic viscosity (μ) of 4 cP which is within the physiological range [59]; L was the length of the capillary; and R was the radius of the capillary. Vessels greater than 10 μm in diameter were excluded from the efficiency calculation as we wanted to examine blood flow through the capillary nexus between arteries and veins where blood flow may be reversible [60]. The sum of the resistivities along the least resistive path specified ⍴ij in the efficiency calculation: this sum quantifies how hard it is for blood to move through the said microvascular bed. N refers to the number of nodes in the network. The efficiency was calculated as the average of the inverse of the least resistive paths between all pairs of nodes. The efficiency increases with increasing radius on the shortest paths through the network.

Removed in the revised version: In computing resistivity, we assumed a fluidic viscosity (μ) of 4 cP which is within the physiological range [59]; L was the length of the capillary; and R was the radius of the capillary. Vessels greater than 10 μm in diameter were excluded from the efficiency calculation as we wanted to examine blood flow through the capillary nexus between arteries and veins where blood flow may be reversible [60]. The sum of the resistivities along the least resistive path specified ⍴ij in the efficiency calculation: this sum quantifies how hard it is for blood to move through the said microvascular bed. N refers to the number of nodes in the network. The efficiency was calculated as the average of the inverse of the least resistive paths between all pairs of nodes. The efficiency increases with increasing radius on the shortest paths through the network.

Statistical models

Earlier version: To statistically compare deep learning model performance, we used the Wilcoxon signed-rank test as implemented in scipy (1.9.3) [61]. Statistics for vascular radius and network metrics were performed in R (4.3.1) using restricted maximum likelihood mixed effects models as implemented in lme4 (1.134) [62], with post hoc comparisons done using emmeans (1.8.8) [63]. We ran the following linear mixed effects model, separately on dilating and constricting vessels that exhibited radius changes above two standard deviations of the vessel’s baseline radius, to examine the effects of photostimulation power on microvascular radius changes in responding vessels: Due to the difference in their vessel wall ultrastructure, larger microvessels (above 5 μm in radius) were examined separately from capillaries (radius < 5 μm).

Revised version: To statistically compare deep learning model performance, we used the Wilcoxon signed-rank test as implemented in scipy (1.9.3) [61]. Statistics for vascular radius and network metrics were performed in R (4.3.1) using restricted maximum likelihood mixed effects models as implemented in lme4 (1.134) [62], with post hoc comparisons done using emmeans (1.8.8) [63]. We ran the following linear mixed effects model, separately on dilating and constricting vessels that exhibited radius changes above two standard deviations of the vessel’s baseline radius, to examine the effects of photostimulation power on microvascular radius changes in responding vessels:

Added in the revised version: ΔRadius ∼ Stimulation + (1|Vessel)

Added in the revised version: Due to the difference in their vessel wall ultrastructure, larger microvessels (above 5 μm in radius) were examined separately from capillaries (radius < 5 μm).

Earlier version: For graph metrics, we included nesting of the field of view within each subject as a random effect so as to account for differences in vascular network architecture within an individual. The linear mixed effects models for the graph metrics used were as follows: In Tables and text, all values have been quoted as mean ± standard deviation, unless otherwise specified.

Revised version: For graph metrics, we included nesting of the field of view within each subject as a random effect so as to account for differences in vascular network architecture within an individual. The linear mixed effects models for the graph metrics used were as follows:

Added in the revised version: Assortativity ∼ Stimulation + (1|Subject/Field of View)

Added in the revised version: Δ Efficiency ∼ Stimulation + (1|Subject/Field of View)

Added in the revised version: In Tables and text, all values have been quoted as mean ± standard deviation, unless otherwise specified.

Results

Application of our computational pipeline resulted in robust segmentation of the vasculature and neurons from 4D in situ 2PFM images and rendering of the microvasculature as a graph. The vessel-wise and vertex-wise calibres were tracked across stimulation conditions and related to the cortical depth and the distance to the closest YFP-expressing neuron, mapping the network-level vascular responses to ChR2 activation and revealing the coordination of the microvascular responses following neuronal activation.

Segmentation model results and comparisons

Earlier version: We compared an ensemble of UNETR models, an ensemble of U-Net models, and an ilastik random forest model on a test dataset of 9 images (507×507×250 μm each) from 6 mice. Examples of segmentation masks produced by each of the models are shown in Figure 4. When evaluating model performance, we paid close attention to the smoothness of the surface of the segmentation masks due to the sensitivity of the centerline extraction algorithms to irregularities in the surface of the masks: smoother vascular segmentation masks resulted in fewer falsely identified vessel branches. Ilastik tended to over-segment vessels, i.e. the model returned numerous false positives, having a high recall (0.89±0.19) but low precision (0.37±0.33) (Figure 3, Supplementary Table 3). When comparing the UNETR and U-Net models, we focused on the surface-based mean surface distance and HD95 distance metrics. Since we observed no significant differences in these metrics between the two models, we selected the UNETR model as our final model because it produced more consistent segmentations on visual inspection and showed significantly better performance than ilastik on HD95 for both vessel and neuron segmentation (p < 0.05).

Revised version: We compared an ensemble of UNETR models, an ensemble of U-Net models, and an ilastik random forest model on a test dataset of 9 images (507×507×250 μm each) from 6 mice. Examples of segmentation masks produced by each of the models are shown in Figure 4 and Supplementary Figure 4. When evaluating model performance, we paid close attention to the smoothness of the surface of the segmentation masks due to the sensitivity of the centerline extraction algorithms to irregularities in the surface of the masks: smoother vascular segmentation masks resulted in fewer falsely identified vessel branches. Ilastik tended to over-segment vessels, i.e. the model returned numerous false positives, having a high recall (0.89±0.19) but low precision (0.37±0.33) (Figure 3, Supplementary Table 3). When comparing the UNETR and U-Net models, we focused on the surface-based mean surface distance and HD95 distance metrics. Since we observed no significant differences in these metrics between the two models, we selected the UNETR model as our final model because it produced more consistent segmentations on visual inspection and showed significantly better performance than ilastik on HD95 for both vessel and neuron segmentation (p < 0.05).

Model performance metrics:

The Dice, precision, recall, mean surface distance, and HD95 distance for the vascular (A) and neuron (B) channels. Each model was evaluated on the same test dataset composed of 9 images (250×507×507 μm each) from 6 mice. A Wilcoxon signed-rank test was used to compare the model’s performance on each performance metric for images from the test dataset. * p < 0.05, ** p < 0.005, and *** p < 0.0005. p-values were not adjusted.

Visual model comparison:

A. Raw images of the vascular channel with the neuron channel subtracted to facilitate vessel visualization. The first and last stacks in each row span from the cortical surface to 250 μm below the surface, while the middle stack spans from 250 μm below the surface to 500 μm below the surface. All images were from the test dataset, which was unseen during model training. B. Ground truth segmentation masks for the vasculature were generated by a rater who utilized ilastik-assisted manual segmentation. C. Ilastik predictions generated via a random forest model. D. Binary segmentation masks generated by an ensemble of 3D UNet Models. E. Binary segmentation masks generated by an ensemble of 3D UNETR models.

Vessel extraction improvements via image registration

Rigid registration across all time points from the same field of view improved the ability to trace vascular paths from in situ 2PFM data. Firstly, registration decreased the mean squared error (MSE) between acquisitions from 1306±747 to 0.008±0.003 Signal Units. The number of images acquired per field of view ranged from a total of 2 to 10 depending on how many repeats were able to be acquired. Registration enabled the computation of the union of segmentation masks from all time points. This increased the number of vessel segments identified in each field of view from 241±174 based on a single time point to 412±281 vessel segments per field of view (507×507×250 μm, n=107 fields of view). Taking the union of segmentation masks of an image stack across all time points substantially decreased the incidence of gaps in capillaries, likely arising due to “transient RBC plugs.” The pipeline’s ability to reconstruct the cortical vascular network was thus significantly improved by registering data obtained at different time points.

Validation of pipeline sensitivity to geometric changes

To evaluate the ability of our computational pipeline to detect vessel caliber changes of various magnitudes, we simulated a range of vascular caliber changes and injected various levels of Gaussian noise. Across >100,000 simulations, the fit of the estimated radius following rescaling against the simulated radius had an R2 value of 0.68. Figure 5B presents a heatmap of the estimated radius post-scaling vs. simulated radius, across different vertices of vessel centerlines, highlighting the ability of our pipeline to estimate vascular radii accurately. The addition of Gaussian noise revealed the robustness of the pipeline: radius estimates remained stable with increasing noise levels, until the addition of noise with a standard deviation of over 200 SU (with the intensity of the images ranging from 0 to 1023 SU).

Estimation of simulated radii changes:

A. An image in the plane orthogonal to the local tangent to a capillary with the detected boundary (in blue) and with the estimated radius of 2.28 μm. On the right, this image was resized (upsampling, via bicubic interpolation, by 1.10 times) to simulate dilation. B. The plot shows correspondence between the estimated radius following scaling and the simulated level of scaling. C. An image in the plane orthogonal to the local tangent of a capillary with the detected boundary (in blue) and with the estimated radius of 3.65 μm. On the right, Gaussian noise with a sigma of 205.36 SU was added to the image. D. The estimated % change in the vessel’s radius after the addition of varying levels of Gaussian noise, demonstrating the robustness of the radius estimated to noise.

Added in the revised version: Our boundary detection algorithm successfully estimated the radius of precisely specified fluorescent beads. The bead images had a signal-to-noise ratio of 6.79 ± 0.16 (about 35% higher than our in vivo images): to match their SNR to that of in vivo vessel data, following deconvolution, we added Gaussian noise with a standard deviation of 85 SU to the images, bringing the SNR down to 5.05 ± 0.15. The data processing pipeline was kept unaltered except for the bead segmentation, performed via image thresholding instead of our deep learning model (trained on vessel data). The bead boundary was computed following the same algorithm used on vessel data: i.e., by the average of the minimum intensity gradients computed along 36 radial spokes emanating from the centreline vertex in the orthogonal plane. To demonstrate an averaging-induced decrease in the uncertainty of the bead radius estimates on a scale that is finer than the nominal resolution of the imaging configuration, we tested four averaging levels in 289 beads. Three of these averaging levels were lower than that used on the vessels, and one matched that used on the vessels (36 spokes per orthogonal plane and a minimum of 10 orthogonal planes per vessel). As the amount of averaging increased, the uncertainty on the diameter of the beads decreased, and our estimate of the bead’s diameter converged upon the manufacturer’s Coulter counter-based specifications (7.32 ± 0.27 μm), as tabulated below in Table 1.

Vascular morphology and heterogeneity within and among vessels

Earlier version: Segmentation coupled with graph extraction enabled a detailed characterization of the microvascular network properties. The morphological properties of extracted networks are listed in Table 1, with the probability densities of the vessel length, baseline vessel radius, mean vessel segment depth, and vessel branch point depth shown in Supplementary Figure 4. On the extracted graphs, the vascular radius and distance to labelled neurons were sampled every 1-1.73 μm, enabling detailed analysis of the relationship between the vessel radius change and the proximity to the YFP-expressing neurons. The radius was tracked across different time points, permitting the analysis of the stimulation-induced change in the vascular caliber. To highlight the ability of the pipeline to detect vessels that significantly change their radius after stimulation, Figure 6A shows the standard deviation of the average radii on each vessel segment during baseline frames for three mice. There was a large difference in this standard deviation across various blood vessels, showcasing the model’s ability to reveal baseline variations within each subject. We examined the average change in the vascular radius of each vessel segment after vs. before photostimulation (Figure 6B), with even finer spatial patterns detected by analyzing the vertex-wise radius changes (Figure 6C). The vascular diameter changes were related to the distance from the vessel’s surface to the closest labelled pyramidal neuron at each vertex of the centerline (Figure 6D). The vertex-wise radii estimation allowed the assessment of the variations in radii changes within individual blood vessels (Figure 7). Notably, capillary radius varied along the vessel length across the baseline frames, by 24 ± 28% of the mean resting radius. Consequently, point measurements in vessel calibers - that are widely reported in the literature - do not permit accurate estimation of the microvessel volume changes. Together, the within- and across-vessel radii estimations illustrate the pipeline’s ability to capture spatial variations in the vascular reactivity and relate it to other morphological features (e.g., the distance to the closest labelled neuron).

Revised version: Segmentation coupled with graph extraction enabled a detailed characterization of the microvascular network properties. The morphological properties of extracted networks are listed in Table 2, with the probability densities of the vessel length, baseline vessel radius, mean vessel segment depth, and vessel branch point depth shown in Supplementary Figure 5. On the extracted graphs, the vascular radius and distance to labelled neurons were sampled every 1-1.73 μm, enabling detailed analysis of the relationship between the vessel radius change and the proximity to the YFP-expressing neurons. The radius was tracked across different time points, permitting the analysis of the stimulation-induced change in the vascular caliber. To highlight the ability of the pipeline to detect vessels that significantly change their radius after stimulation, Figure 6A shows the standard deviation of the average radii on each vessel segment during baseline frames for three mice. There was a large difference in this standard deviation across various blood vessels, showcasing the model’s ability to reveal baseline variations within each subject. We examined the average change in the vascular radius of each vessel segment after vs. before photostimulation (Figure 6B), with even finer spatial patterns detected by analyzing the vertex-wise radius changes (Figure 6C). The vascular diameter changes were related to the distance from the vessel’s surface to the closest labelled pyramidal neuron at each vertex of the centerline (Figure 6D). The vertex-wise radii estimation allowed the assessment of the variations in radii changes within individual blood vessels (Figure 7). Notably, capillary radius varied along the vessel length across the baseline frames, by 24 ± 28% of the mean resting radius. Consequently, point measurements in vessel calibers - that are widely reported in the literature - do not permit accurate estimation of the microvessel volume changes. Together, the within- and across-vessel radii estimations illustrate the pipeline’s ability to capture spatial variations in the vascular reactivity and relate it to other morphological features (e.g., the distance to the closest labelled neuron).

Vascular graph examples:

A. Baseline variability in vessel diameter estimated by the standard deviation of each vessel’s mean radius across baseline time frames. B. Mean change in the vessel radius induced by optogenetic stimulation. C. Mean change in the vertexwise radius, allowing the visualization of heterogeneity of radius changes within each vessel. D. Distance from each vertex to the closest pyramidal neuron. Each row corresponds to the vascular graph of a different mouse.

Vertex-wise radii along vessel lengths of a sample artery, capillary and venule at baseline vs. post-stimulation:

A. MIP of an artery, vein, and capillary segments before (left) and after (right) optogenetic stimulation with 458nm light at 1.1 mW/mm2. The artery and capillary dilated by 1.33±0.86 μm and 0.42±0.39 μm, respectively (for both p<1e-4, Mann-Whitney U test), whereas there was no significant change in the venular caliber upon photostimulation (p=0.22, Mann-Whitney U test). B. Estimates of the vertex-wise radius obtained along each of the three vessels’ centrelines, before and after stimulation. C. Vertex-wise radii changes in response to optogenetic stimulation. D. The vertex-wise distance from the vascular surface to the closest YFP-expressing neuron.

Vascular reactivity to optogenetic stimulation

Earlier version: The ability of the pipeline to reveal novel spatial relationships between the vascular network reactivity and labelled neurons was demonstrated by examining the relationship between photostimulation-induced microvascular radii responses and 1) the closest YFP-labelled pyramidal neurons within 80μm, and 2) the cortical depth, at the vertex-wise level and across different photostimulations (Figure 8). Vessels were coarsely segregated into small (average radius < 5 μm) and large (average radius > 5 μm) vessels, as we expected them to respond differently due to their differential wall-associated cell composition [6], [64], [65], [66], [67], [68], [69], [70], [71]. Only vessels longer than 20 μm (i.e. vessels whose radius was computed by averaging over many cross-sectional planes) that significantly responded following optogenetic stimulation were analyzed: a vessel was deemed a responder if its radius changed by more than twice the baseline standard deviation in the vessel’s radius. The morphometric properties of the responders, under different stimulation conditions, are listed in Table 2. The average magnitude of significant microvascular radius changes across all stimulation conditions was 1.04 ± 1.11 μm (62±47%). The variability in the radius change within the vessel was higher in the dilating vessels, 0.77 ± 0.61 μm (66±72%) than in the constricting vessels, 0.69 ± 0.49 μm (46±18%), for 458-nm, photostimulation (p<1e-4) (with no statistically significant changes detected for either 458-nm, photostimulation or 552-nm, photostimulation).

Revised version: The ability of the pipeline to reveal novel spatial relationships between the vascular network reactivity and labelled neurons was demonstrated by examining the relationship between photostimulation-induced microvascular radii responses and 1) the closest YFP-labelled pyramidal neurons within 80μm, and 2) the cortical depth, at the vertex-wise level and across different photostimulations (Figure 8). Vessels were coarsely segregated into small (average radius < 5 μm) and large (average radius > 5 μm) vessels, as we expected them to respond differently due to their differential wall-associated cell composition [6], [64], [65], [66], [67], [68], [69], [70], [71]. Only vessels longer than 20 μm (i.e. vessels whose radius was computed by averaging over many cross-sectional planes) that significantly responded following optogenetic stimulation were analyzed: a vessel was deemed a responder if its radius changed by more than twice the baseline standard deviation in the vessel’s radius. The morphometric properties of the responders, under different stimulation conditions, are listed in Table 3. The average magnitude of significant microvascular radius changes across all stimulation conditions was 1.04 ± 1.11 μm (62±47%). The variability in the radius change within the vessel was higher in the dilating vessels, 0.77 ± 0.61 μm (66±72%) than in the constricting vessels, 0.69 ± 0.49 μm (46±18%), for 458-nm, photostimulation (p<1e-4) (with no statistically significant changes detected for either 458-nm, photostimulation or 552-nm, photostimulation).

Optogenetic activation-induced changes in vessel-wise microvascular radii:

Capillary responses included both dilatations, shown in A. and constrictions, shown in B., with changes in the magnitude of the capillary response with increased photostimulation power. * p < 0.05, ** p < 0.005, and *** p < 0.0005. p-values were not adjusted. C. Probability density function of constrictions and dilations for the 4.3 mW/mm2 photostimulation. D. Changes to capillary radii are displayed in relation to the closest pyramidal neurons. The proportion of vessels constricting increased with the higher intensity of blue light stimulation, and constrictions tended to occur further away from pyramidal neurons than did dilations. E. Mean cortical depth of responding capillaries showed a tendency for dilators to be closer to the surface and for constrictors to be deeper in the tissue.

Added in the revised version: Microvascular Network Coordination Following Optogenetic Stimulation:

Added in the revised version: A. Graph representation of a vascular network of 425 vascular segments from a single image stack. Vessel segments are depicted as nodes of the graph; vascular segments that are joined at junctions are connected by edges. Nodes are coloured by the change in the mean vessel-wise radius following photostimulation with 458 nm light at 4.3 mW/mm2. B. Assortativity of photostimulation-induced changes in mean capillary radius increased with increasing photostimulation power. C. Photostimulation-induced changes in the efficiency of the capillary network. The capillary network efficiency changed by a median -0.16 PΩ-1 (IQR: -0.39 to 0.10 PΩ-1) in response to green light; -0.14 PΩ-1 (IQR: -0.55 to 0.27 PΩ-1) in response to lower intensity bluelight; and 0.22 PΩ-1 (IQR = -0.43;1.47 PΩ-1) in response to higher intensity blue light.. There was a significant increase (p = 0.03) in the capillary network efficiency post 458-nm light at 4.3 mW/mm2, when compared to that following the control green illumination. The measurements came from 72 paired acquisitions of 32 image stacks acquired in 17 mice (9M/8F). * p < 0.05, ** p < 0.005, and *** p < 0.0005. p-values were not adjusted.

Excluding vessels that did not change their radius by over twice the baseline standard deviation removed almost all large vessels from this analysis. (Notwithstanding, Supplementary Figure 5 depicts unfiltered large vessel constrictions and dilations.) We ran mixed effects models (at the vessel level) separately on constricting and dilating vessels to investigate the effect of optogenetic stimulation power on the vessel radius changes. Each of the models was run separately on small and large vessels.

Vessels further away from labelled neurons constrict while those closer to the activated neurons dilate

Removed in the revised version: We examined the relationship between vascular radius changes and the distance to the closest labelled pyramidal neuron, as microvascular response is thought to result from neuronal activation-elicited generation of vasoactive molecules that diffuse to the neighbouring vessels.

Earlier version: For the control condition (552nm, photostimulation), 2.9% of small capillaries dilated while 1.0% of small capillaries constricted; in larger vessels, barely any responded (0.4% dilated). For this control condition, there was no significant difference in the distance from constrictors or dilators to the closest pyramidal neuron. For the 458 nm photostimulation, capillary constrictors were on average farther away than were dilators from the labelledlabelled pyramidal neuron: dilations occurred 16.8 ± 13.5 μm away from labelledlabelled neurons while constrictions occurred 22.7 ± 16.3 μm for photostimulation (p=1.5e-3) whereas the photostimulation had dilations occur 16.1 ± 14.3 μm away and 21.9 ± 14.6 μm for constrictors (p<1e-4). There was no significant shift between the distance from vessels to neurons for the and stimulations with 458 nm light. Dilations in capillaries following 458 nm photostimulation were larger than those following the 552 nm control photostimulation: 0.90 ± 0.93 μm dilatations occurred with and 0.90 ± 0.78 μm with at 458 nm; vs. 0.58 ± 0.92 μm with at 552 nm (p<1e-4). For constrictions, 458 nm photostimulations led to -1.39 ± 1.51 μm radius changes with and -1.20 ± 1.13 μm radius changes with (p=4.4e-3); whereas 552 nm photostimulation induced -0.37 ± 0.30 μm radius changes with of power, which was smaller than the 458-nm induced responses (p=0.02).

Revised version: We examined the relationship between vascular radius changes and the distance to the closest labelled pyramidal neuron, as microvascular response is thought to result from neuronal activation-elicited generation of vasoactive molecules that diffuse to the neighbouring vessels. For the control condition (552nm, photostimulation), 2.9% of small capillaries dilated while 1.0% of small capillaries constricted; in larger vessels, barely any responded (0.4% dilated). For this control condition, there was no significant difference in the distance from constrictors or dilators to the closest pyramidal neuron. For the 458 nm photostimulation, capillary constrictors were on average farther away than were dilators from the labelledlabelled pyramidal neuron: dilations occurred 16.8 ± 13.5 μm away from labelledlabelled neurons while constrictions occurred 22.7 ± 16.3 μm for photostimulation (p=1.5e-3) whereas the photostimulation had dilations occur 16.1 ± 14.3 μm away and 21.9 ± 14.6 μm for constrictors (p<1e-4). There was no significant shift between the distance from vessels to neurons for the and stimulations with 458 nm light.. Dilations in capillaries following 458 nm photostimulation were larger than those following the 552 nm control photostimulation: 0.90 ± 0.93 μm dilatations occurred with and 0.90 ± 0.78 μm with at 458 nm; vs. 0.58 ± 0.92 μm with at 552 nm (p<1e-4). For constrictions, 458 nm photostimulations led to -1.39 ± 1.51 μm radius changes with and -1.20 ± 1.13 μm radius changes with (p=4.4e-3); whereas 552 nm photostimulation induced -0.37 ± 0.30 μm radius changes with of power, which was smaller than the 458-nm induced responses (p=0.02).

Vascular radius changes at increasing cortical depths

Earlier version: Vascular responses were next segregated by the cortical depth of the vessel (i.e. the average vessel distance from the cortical surface) [72]. Dilators tended to be located closer to the cortical surface across all stimulation conditions. Constricting vessels were located at an average 58 ± 187 μm deeper than dilators for 458-nm stimulation at and 37 ± 179 μm deeper for 458-nm photostimulation at (with no change in the mean depth of either constricting or dilating vessels with changes in the photostimulation power).

Revised version: Vascular responses were next segregated by the cortical depth of the vessel (i.e. the average vessel distance from the cortical surface) [72]. Dilators tended to be located closer to the cortical surface across all stimulation conditions. Constricting vessels were located at an average 58 ± 187 μm deeper than dilators for 458-nm stimulation at (p=0.02), and 37 ± 179 μm deeper for 458-nm photostimulation at (p<1e-4) (with no change in the mean depth of either constricting or dilating vessels with changes in the photostimulation power).

Vascular Network Coordination Following Optogenetic Stimulation

We examined the coordination of changes in the microvascular network as a whole via assortativity of radius changes and network efficiency changes. The vessel responses were observed to be assortative, i.e., capillaries mirrored the responses in their neighbours. The increases in stimulation power were accompanied by increases in the assortativity of capillary responses: increasing stimulation level resulted in heightened coordination between adjacent capillaries.

The efficiency increased only at the strongest blue photostimulation, i.e., only at this level of stimulation did the resistivity along the average of all of the shortest paths between junctions in the vascular network decrease, resulting in attenuated resistance to flow through the network. The distribution of changes to the efficiency were highly skewed (with a coefficient of skewness of -1.06 for green illumination, 2.92 for lower intensity blue photostimulation, and 4.87 for higher intensity blue photostimulation). The median increase in the efficiency induced by the higher intensity blue photostimulation, of 4% (IQR: -8% to 38%), was significantly higher than the median -6% (IQR=-9% to 4%) efficiency change following the control green illumination

Discussion

Recent studies have demonstrated temporal propagation and coordination in cerebrovascular responses to neuronal activation, whereby arteries dilated after capillary exposure to increased potassium ion concentration [73]; opposing geometric changes have been reported by some studies in capillaries connected by intercapillary tunnelling nanotubes [24]. However, how the effects of these and other mechanisms influence in situ reactivity of the 3D brain capillary network remains unknown. Here, we developed a pipeline for extracting graphs of brain microvascular networks from in situ 2PFM and examine coordination within and across capillaries. Capillary networks and their geometrical changes were imaged via 2PFM during periods of baseline alternated with photostimulation of ChR2 in pyramidal neurons of transgenic mice, and the microvascular mesh was evaluated every 1-1.73 μm. The vascular morphology was then analyzed vertex- and vessel-wise across the entire network. All vessels exhibited significant heterogeneity in caliber changes along their length. Neuronal activation induced both dilations and constrictions of vessels and the incidence of constrictions increased with increasing cortical depth. As the stimulation power increased, the tendency for vessels to change their radius by an amount similar to their neighbours increased. Only the highest photostimulation intensity elicited an increase in the network efficiency. Our findings reveal an intricate level of coordination among brain microvessels and provide a computational analysis platform for interrogating a host of hypotheses on cerebral microvascular reactivity.

Vascular Segmentation and Network Extraction

Earlier version: Intensity thresholding-based image processing pipelines have been used to examine vascular networks and quantify vascular morphology, but they have not gained widespread use due to difficulties in adapting them to highly heterogeneous levels of noise across samples [68], [74], [75], [76]. Deep learning-based methods for analyzing vascular morphology from 3D microscopy images have become prevalent as they provide robust segmentation results over a wide range of signal-to-noise ratios. Recent work has demonstrated steady improvements of segmentation models’ performance with respect to similarity-based metrics (i.e., Dice scores) [77], [78], [79], [80], [81], [82]; though surface-based metrics may be better predictors of how amenable segmentation outputs will be to subsequent morphological analysis of the microvascular network. Our final UNETR model was selected based on a combination of performance metrics including mean surface distance, Hausdorff 95% distance, and rater evaluation, to maximize the smoothness of the surface of generated segmentation masks and reduce false positive branch points during centerline extraction; thereby leading to higher fidelity rendering of microvascular networks. Graph generation was greatly facilitated by computing the union of the vascular segmentation masks across all time points as it enabled tracing of capillaries that had stalls at the individual time points (since the accompanying loss of the fluorescent label otherwise resulted in graph discontinuities).

Revised version: Intensity thresholding-based image processing pipelines have been used to examine vascular networks and quantify vascular morphology, but they have not gained widespread use due to difficulties in adapting them to highly heterogeneous levels of noise across samples [68], [74], [75], [76]. Deep learning-based methods for analyzing vascular morphology from 3D microscopy images have become prevalent as they provide robust segmentation results over a wide range of signal-to-noise ratios. Recent work has demonstrated steady improvements of segmentation models’ performance with respect to similarity-based metrics (i.e., Dice scores) [77], [78], [79], [80], [81], [82]; though surface-based metrics may be better predictors of how amenable segmentation outputs will be to subsequent morphological analysis of the microvascular network. Our final UNETR model was selected based on a combination of performance metrics including mean surface distance, Hausdorf 95% distance, and rater evaluation, to maximize the smoothness of the surface of generated segmentation masks and reduce false positive branch points during centerline extraction; thereby leading to higher fidelity rendering of microvascular networks. Graph generation was greatly facilitated by computing the union of the vascular segmentation masks across all time points as it enabled tracing of capillaries that had stalls at the individual time points (since the accompanying loss of the fluorescent label otherwise resulted in graph discontinuities).

Earlier version: We tested the ability of the pipeline to detect changes to simulated changes to images and the pipeline’s sensitivity to perturbations to extracted vascular morphology. Using resizing, we confirmed that the boundary detection algorithm was able to detect prescribed changes (cf. Figure 5). The average radius estimates for vessels were then shown to vary by 0.64±3.44% upon changes in centerline position, demonstrating that the radius estimates have a low sensitivity to small (2-3 μm) perturbations in centerline placement. Visual examples of vessels repeatedly dilating or constricting are shown in Supplementary Figures 6 and 7.

Revised version: We tested the ability of the pipeline to detect changes to simulated changes to images and the pipeline’s sensitivity to perturbations to extracted vascular morphology. Using resizing, we confirmed that the boundary detection algorithm was able to detect prescribed changes (cf. Figure 5). The average radius estimates for vessels were then shown to vary by 0.64±3.44% upon changes in centerline position, demonstrating that the radius estimates have a low sensitivity to small (2-3 μm) perturbations in centerline placement. Visual examples of vessels repeatedly dilating or constricting are shown in Supplementary Figures 7 and 8.

Network morphological properties at baseline were in line with prior work. The vascular length density measured from fixed tissue ranged from 0.44 to 1.10 m/mm3 [74], [83], [84], [85], [86], [87]. Our reported vascular density in the forelimb region of the primary somatosensory cortex was 0.40 ± 0.22 m/mm3, with the low end of the range value expected due to fluorescence absorption by hemoglobin in the large pial vessels leading to signal dropout (or shadowing) in the underlying tissue. Our reported average capillary radius of 2.19 ± 1.66 μm was also in line with other studies, where the mean capillary radius ranged from 1.75 to 2.2 μm as measured with confocal microscopy or 2PFM [7], [74].

Earlier version: In Situ Changes to Vessel Calibers Upon Neuronal Activation

Revised version: Next In Situ Changes to Vessel Calibers Upon Neuronal Activation

Caliber changes at individual vertices along vessel centerlines exhibited significant heterogeneity. Such heterogeneity is expected due to non-uniformly distributed alpha smooth muscle actin-containing cells along vessel walls, as well as differential activations leading to heterogeneous metabolic demand within the tissue [1], [6], [88], [89], [90], [91]. Many previous studies assumed vessel caliber changes to be uniform, compromising the accuracy of the estimates.

Earlier version: As expected, the control 552 nm stimulation led to minimal changes in vessel calibers. To probe for off-target effects, non-transgenic mice were also tested with the same optical setup and photostimulation, with no changes to vascular diameters observed at any of the photostimulation powers utilized (cf. Supplementary Figure 8). In transgenic mice, we detected an average capillary dilation in significantly responding vessels of 70±83% with low-intensity 458 nm stimulation, and 67±61% with higher intensity 458 nm stimulation. Across photostimulation conditions, the capillary dilations ranged from 2% to 805%. These calibre changes were higher than those previously reported, which varied from 2% to 20% depending on the capillary branch order [6], [7], [23], [92], [93]. Far less data are available on constrictions. In the current work, constrictions averaged 47±20% for lower-intensity blue light stimulation and 47±17% for higher-intensity blue light stimulation, with a constriction range of 5% to 97% of the baseline radius. These are higher than the previously reported constrictions of 20% [6], [23], likely due to our identifying as responding vessels only those whose caliber changed by at least twice their baseline caliber variation. It is also worth noting that vessels’ response directions were consistent on repeated trials. Of the vessels whose radius change exceeded twice the baseline variability across time, 31.7% dilated on some trials while constricting on others; 41.1% dilated on each trial; and 27.2% constricted on each trial. (Note that some trials use 1.1 vs. 4.3 mW/mm2 and some have opposite scanning directions).

Revised version: As expected, the control 552 nm stimulation led to minimal changes in vessel calibers. To probe for off-target effects, non-transgenic mice were also tested with the same optical setup and photostimulation, with no changes to vascular diameters observed at any of the photostimulation powers utilized (cf. Supplementary Figure 9). In transgenic mice, we detected an average capillary dilation in significantly responding vessels of 70±83% with low-intensity 458 nm stimulation, and 67±61% with higher intensity 458 nm stimulation. Across photostimulation conditions, the capillary dilations ranged from 2% to 805%. These calibre changes were higher than those previously reported, which varied from 2% to 20% depending on the capillary branch order [6], [7], [23], [92], [93]. Far less data are available on constrictions. In the current work, constrictions averaged 47±20% for lower-intensity blue light stimulation and 47±17% for higher-intensity blue light stimulation, with a constriction range of 5% to 97% of the baseline radius. These are higher than the previously reported constrictions of 20% [6], [23], likely due to our identifying as responding vessels only those whose caliber changed by at least twice their baseline caliber variation. It is also worth noting that vessels’ response directions were consistent on repeated trials. Of the vessels whose radius change exceeded twice the baseline variability across time, 31.7% dilated on some trials while constricting on others; 41.1% dilated on each trial; and 27.2% constricted on each trial. (Note that some trials use 1.1 vs. 4.3 mW/mm2 and some have opposite scanning directions).

458-nm photostimulation resulted in a mix of constrictions and dilations with 44.1% of significantly responding vessels within 10 μm of a labelled pyramidal neuron constricting and 55.1% dilating, while 53.3% of vessels further than 30 μm constricted and 46.7% dilated. The cutoff distances from the closest labelled neuron were based on estimates of cerebral metabolic rate of oxygen consumption that showed a steep gradient in oxygen consumption with distance from arteries, CMRO2 being halved by 30 μm away [94]. The stronger blue light stimulation led to an increased rate of constrictions, double that of the low-powered blue light stimulation. For larger vessels, both 458 nm stimulation powers led to a similar dilation level that diminished with increasing distance from labelled pyramidal neurons. This tendency for vessels close to neurons to dilate and further away ones to constrict would be expected in flow redirection into regions of high level of neuronal activity. Stimulation power dependence in blood flow changes has previously been reported in optogenetic mouse models with diffuse stimulation via LED probes, and following transcranial alternating current stimulation [95], [96]. However, neither of the previously employed methods was able to discern the spatial relationship between the vascular caliber changes, or relate these changes to the distribution of the stimulated neurons.

As the blue light stimulation power increased, the mean depth of both constricting and dilating vessels increased, likely resulting from higher intensity light reaching pyramidal neurons deeper in the tissue [97], [98]. The blue light would be expected to excite a lower number of neurons farther from the cortical surface at lower powers. Our results underscore that the haemodynamic response following targeted neuronal activation is not uniformly distributed across the microvascular network: accurate neurovascular coupling assessment thus requires network-based analysis.

Vascular Network Reactivity

To study the microvascular network response as a whole, we examined the assortativity between capillary radius changes and network efficiency changes following optogenetic stimulation. These two graph theory metrics were selected as they both leverage the knowledge of the vascular network structure. Assortativity sheds light on how the vascular network coordinates its responses, while efficiency provides insight into the extent to which those changes facilitate flow through the network. The assortativity revealed that as the stimulation power increased, the tendency of vessels to match their changes to those of their neighbours increased. Previously characterized assortative mechanisms include endothelial cell cation conduction via Kir2.1 channels to synchronize vascular responses [73], and spatial adjacency of pericytes on in vitro retinal preparation leading to assortative changes in neighbouring capillaries [88]. Disassortative (causing opposite changes) mechanisms of capillary coordination have also previously been observed in situ and may result from intercapillary nanotubules’ signalling causing connected pericytes to undergo opposing changes [24]. While not ruling out the presence of disassortative control mechanisms, our results suggest that assortative mechanisms dominate capillary responses to neuronal activation in the somatosensory cortex.

Earlier version: The network efficiency here can be thought of as paralleling mean transit time, i.e., the time it takes blood to traverse the capillary network from the arteries to the veins. In situ studies of mean transit time have revealed a high heterogeneity of plasma traversal of the capillary bed during stimulation, with stimulation reducing plasma transit times by 11% to 20% from its resting levels [93], [99], and simulations suggesting that capillary network geometry and locations of caliber changes exert a substantial influence on these responses [100]. The efficiency of the vascular network here increased significantly only with the strongest 458 nm stimulation. Small dilatations may thus not increase flow in the cortex. The differences in efficiency are likely due to the patterns of localized dilations and constrictions within the vascular network. Efficiency calculations are sensitive to bottlenecks when traversing meshes and certain locations constricting or dilating can have profound impacts on the shortest paths between nodes and the path’s resistivity. The highest powered 458 nm stimulation increasing efficiency may have resulted from increased assortativity causing dilation in key locations within the microvascular network, leading to a significant reduction in shortest path resistivity.

Revised version: The network efficiency here can be thought of as paralleling mean transit time, i.e., the time it takes blood to traverse the capillary network from the arteries to the veins. In situ studies of mean transit time have revealed a high heterogeneity of plasma traversal of the capillary bed during stimulation, with stimulation reducing plasma transit times by 11% to 20% from its resting levels [93], [99], and simulations suggesting that capillary network geometry and locations of caliber changes exert a substantial influence on these responses [100]. The efficiency of the vascular network here increased significantly only with the strongest 458 nm stimulation. Small dilatations may thus not increase flow in the cortex. The differences in efficiency are likely due to the patterns of localized dilations and constrictions within the vascular network.

Added in the revised version: Efficiency calculations are sensitive to bottlenecks when traversing meshes and certain locations constricting or dilating can have profound impacts on the shortest paths between nodes and the path’s resistivity. The highest powered 458 nm stimulation increasing efficiency may have resulted from increased assortativity causing dilation in key locations within the microvascular network, leading to a significant reduction in shortest path resistivity.

Added in the revised version: Comparison with commercial and open-source vascular analysis pipelines

Added in the revised version: To compare our results with those achievable on these data with other pipelines for segmentation and graph network extraction, we compared segmentation results qualitatively with Imaris version 9.2.1 (Bitplane) and vascular graph extraction with VesselVio [101]. For the Imaris comparison, three small volumes were annotated by hand to label vessels. Example slices of the segmentation results are shown in Supplementary Figure 10. Imaris tended to either over- or under-segment vessels, disregard fine details of the vascular boundaries, and produce jagged edges in the vascular segmentation masks. In addition to these issues with segmentation mask quality, manual segmentation of a single volume took days for a rater to annotate. To compare to VesselVio, binary segmentation masks (one before and one after photostimulation) generated with our deep learning models were loaded into VesselVio for graph extraction, as VesselVio does not have its own method for generating segmentation masks. This also facilitates a direct comparison of the benefits of our graph extraction pipeline to VesselVio. Visualizations of the two graphs are shown in Supplementary Figure 11. Vesselvio produced many hairs at both time points, and the total number of segments varied considerably between the two sequential stacks: while the baseline scan resulted in 546 vessel segments, the second scan had 642 vessel segments. These discrepancies are difficult to resolve in post-processing and preclude a direct comparison of individual vessel segments across time. As the segmentation masks we used in graph extraction derive from the union of multiple time points, we could better trace the vasculature and identify more connections in our extracted graph. Furthermore, VesselVio relies on the distance transform of the user supplied segmentation mask to estimate vascular radii; consequently, these estimates are highly susceptible to variations in the input segmentation masks.We repeatedly saw slight variations between boundary placements of all of the models we utilized (ilastik, UNet, and UNETR) and those produced by raters. Our pipeline mitigates this segmentation method bias by using intensity gradient-based boundary detection from centerlines in the image (as opposed to using the distance transform of the segmentation mask, as in VesselVio).

Pipeline Limitations and Adaptability

Earlier version: The segmentation model was trained only on vascular and neuronal labels, limiting its generalizability to segmenting alternative cells in the current state. However, it can easily be fine-tuned or retrained to label other brain cells (e.g., pericytes, astrocytes, or endothelial cells). Our vascular segmentation model generalized well to C57BL/6 mouse and Fischer rat data, as well as to Thy1-ChR2 light-sheet fluorescence microscopy images gathered on an UltraMicroscope Blaze lightsheet fluorescence microscope (Miltenyi Biotech) (cf. Supplementary Figure 9 and Supplementary Table 3). However, the segmentation model performed poorly when significant bleeding occurred in the cranial window, compromising the vascular contrast. Our imaging protocol was challenged by the need to resolve individual vessel responses and capture the entire network within the span of the microvascular response to stimulation: thus, we could not examine the kinetics of the microvascular response. The temporal evolution of the response in individual vessels, however, has been reported on using line scanning acquisitions to measure red blood cell velocity and flux [6], [9], [23], [30], [32], [93]. Moreover, the present analysis pipeline can readily be applied to 2PFM data obtained with finer temporal (e.g., via a Piezo objective positioner) or spatial resolution, and/or different size fields of view. Additionally, alternative definitions of responding vessels may be useful depending on the end goal of a study (e.g. selecting a threshold for the radius change based on a percentage change from the baseline level: cf. Supplementary Figure 10 for capillary changes above 10% of the baseline radius). Finally, microvascular networks in different brain areas may show distinct spatiotemporal profiles of response to neuronal activation. Future work is required to test the generalizability of present findings across different brain regions.

Revised version: The segmentation model was trained only on vascular and neuronal labels, limiting its generalizability to segmenting alternative cells in the current state. However, it can easily be fine-tuned or retrained to label other brain cells (e.g., pericytes, astrocytes, or endothelial cells). Our vascular segmentation model generalized well to C57BL/6 mouse and Fischer rat data, as well as to Thy1-ChR2 light-sheet fluorescence microscopy images gathered on an UltraMicroscope Blaze lightsheet fluorescence microscope (Miltenyi Biotech) (cf. Supplementary Figure 12, 13 and Supplementary Table 3). However, the segmentation model performed poorly when significant bleeding occurred in the cranial window, compromising the vascular contrast. Our imaging protocol was challenged by the need to resolve individual vessel responses and capture the entire network within the span of the microvascular response to stimulation: thus, we could not examine the kinetics of the microvascular response. The temporal evolution of the response in individual vessels, however, has been reported on using line scanning acquisitions to measure red blood cell velocity and flux [6], [9], [23], [30], [32], [93]. Moreover, the present analysis pipeline can readily be applied to 2PFM data obtained with finer temporal (e.g., via a Piezo objective positioner) or spatial resolution, and/or different size fields of view. Additionally, alternative definitions of responding vessels may be useful depending on the end goal of a study (e.g. selecting a threshold for the radius change based on a percentage change from the baseline level: cf. Supplementary Figure 14 for capillary changes above 10% of the baseline radius). Finally, microvascular networks in different brain areas may show distinct spatiotemporal profiles of response to neuronal activation. Future work is required to test the generalizability of present findings across different brain regions.

Conclusion

We developed a novel deep learning-based computational pipeline for analysis of a time series of 3D 2PFM images and investigation of spatial patterns in microvascular network reactivity to neuronal activation. The microvascular network was represented as a graph, allowing for the evaluation of network geometry changes over time. We tracked the size of blood vessels throughout the network and related vessel radius changes to the distance from the stimulated neurons and the cortical depth. Neuronal activation induced both dilatations and constrictions of capillaries, and the magnitude of these responses increased with increased photostimulation levels while showing significant heterogeneity within and between vessels. In the analysis presented, vertex-wise measurements were aggregated for vessel-wise analysis resulting in highly robust estimates of vessels’ calibers and allowing ready comparisons to literature. Notwithstanding, the pipeline also affords vertex-wise analysis and thus registration of microvascular reactivity with other local morphological features, at an unprecedented spatial scale. With increasing distance of the vessel from the most proximal activated neuron, dilatation magnitude decreased and the incidence of constrictions increased. At the highest stimulation level investigated, the incidence of vessel constrictions also increased with cortical depth. With increasing activation levels, capillaries displayed diameter changes that were similar to their immediate neighbours, while vascular network efficiency increased only under the strongest stimulation. Our computational analysis pipeline permits probing microvascular network reactivity and sheds light on the heterogeneity and coordination of vessel caliber changes across the microvascular network.The pipeline will be made available to the research community to propel future studies of neurovascular coupling and network reactivity.